Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 18
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 12<sup>th</sup> 2019, 5 pm**

# Trie

A `trie` is a kind of search tree. All the descendants of a node have a **common prefix** of the string associated with that node, and the root is associated with the empty string.

The `trie` below represents the following names: 'John', 'Johnny', 'Joe', 'Jane', 'Jack'. Notice that 'John' and 'Johnny' share a common prefix: 'John'. 'John', 'Johnny', and 'Joe' share a common prefix: 'Jo' and all the names share 'J' as a common prefix. This allows us to efficiently store these names.

<img src="trie0.jpg" width="150">

This kind of a data structure will be helpful for checking to see if a word is in a dictionary. It might also, given a prefix, be used to generate a list of words that complete that prefix.

# Problem 1

Implement the `trie_node` class. This should include a single-character string `val`, a dictionary `children`, and a boolean value `is_complete` that is `True` when the node is the last character in a valid word.

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
class trie_node(object):
    def __init__(self, val):
        self.val=val
        self.children={}
        self.is_complete=False
```

```python
a=trie_node('a')
b=trie_node('b')
a.children['b']=b
a.children
a.is_complete
```

Output:
```
False
```

<!-- grade_id: cell-b1b12dbb4976108d, points: 0.5, grade: true, solution: false -->
```python
root = trie_node('')
root.children['m'] = trie_node('m')
root.children['n'] = trie_node('n')

assert root.val == ''
assert root.children['m'].val == 'm'
assert root.children['n'].val == 'n'
assert len(root.children) == 2
```

# Problem 2

Implement the `trie` class. This should include a single `root` trie_node.

You must implement the following methods:

`find`: takes a string `word` as input and returns `True` if `word` is in the `trie` and `False` otherwise.

`insert`: takes a string `word` as input and inserts `word` into the `trie`.

`starts_with`: takes a string `prefix` as input and returns a list of all strings in the `trie` that have `prefix` as a prefix.

<!-- grade_id: cell-26254032e02425bf, points: 0, grade: false, solution: true -->
```python
class trie(object):
    def __init__(self):
        self.root = trie_node('')
        
    def __str__(self):
        result = ''
        q = [self.root]
        
        while len(q) > 0:
            node = q[0]
            q = q[1:]
            
            result += node.val + ' '
            
            if len(node.children) > 0:
                for v in node.children.values():
                    q.append(v)
        
        return result

    def insert(self, word):
        index=0
        parent=self.root
        while index<len(word):
            newT=trie_node(word[index])
            contents=parent.children.keys()
            if newT.val in contents:
                parent=parent.children[word[index]]
                index+=1
                continue
            else:
                parent.children[word[index]]=newT
                parent=parent.children[word[index]]
                index+=1
                continue
        parent.is_complete=True
                  
    def find(self, word):
        
        def helper(word,index,parent):
            if index==len(word):
                if parent.is_complete==True:
                    return True
                else:
                    return False
            else: 
                if word[index] not in parent.children:
                    return False
                else:
                    return helper(word,index+1,parent.children[word[index]])
            
        return helper(word,0,self.root)
        
    def starts_with(self,prefix):
        R=[[self.root]]
        store=R[:]
        R=[]
        done=[]
        while store!=[]:
            for i in store:
                if i[-1].is_complete==True:
                    done.append(i)
                if i[-1].children!={}:
                    for j in i[-1].children.values():
                        R.append(i+[j])
            store=R[:]
            R=[]
        result=[]
        for List in done:
            for i in range(len(List)):
                List[i]=List[i].val
            List=''.join(List)
            if List.startswith(prefix)==True:
                result.append(List)
        if prefix in result:
            result.remove(prefix)
        return result

```

## Test Case 1: Insert

<!-- grade_id: cell-68689597d1269994, points: 0.5, grade: true, solution: false -->
```python
t = trie()

t.insert('John')
t.insert('Johnny')
t.insert('Joe')
t.insert('Jane')
t.insert('Jack')

assert str(t).strip() == 'J o a h e n c n e k n y'
```

## Test Case 2: Find

<!-- grade_id: cell-6ff65e72731d281c, points: 0.5, grade: true, solution: false -->
```python
t = trie()

J = trie_node('J')
o = trie_node('o')
h = trie_node('h')
n1 = trie_node('n')
n1.is_complete = True 
n2 = trie_node('n')
y = trie_node('y')
y.is_complete = True
a = trie_node('a')
n3 = trie_node('n')
e = trie_node('e')
e.is_complete = True

t.root.children['J'] = J
J.children['o'] = o
J.children['a'] = a
o.children['h'] = h
h.children['n'] = n1
n1.children['n'] = n2
n2.children['y'] = y
a.children['n'] = n3
n3.children['e'] = e

assert t.find('John') == True
assert t.find('Johnn') == False
assert t.find('Johnny') == True
assert t.find('Jane') == True
assert t.find('Janey') == False
```

## Test Case 3: Starts With

<!-- grade_id: cell-b920d7b03d3b969c, points: 0.5, grade: true, solution: false -->
```python
dictionary = trie()

file = open('wordlist.txt', 'r')

for word in file:
    dictionary.insert(word.strip())

assert set(dictionary.starts_with('abac')) == set(['abaca', 'abacist', 'aback', 'abactinal', 'abacus'])
assert set(dictionary.starts_with('chin')) == set(['chin-up', 'china', 'chinaberry', 'chinaman', 'chinaware', 'chincapin', 'chinch', 'chincherinchee', 'chinchilla', 'chinchillidae', 'chine', 'chinese', 'chink', 'chinked', 'chinking', 'chinks', 'chinless', 'chino', 'chinoiserie', 'chinoises', 'chinook', 'chinookan', 'chinos', 'chintz'])
assert set(dictionary.starts_with('xenop')) == set(['xenophobia', 'xenophobic', 'xenopodidae', 'xenopus'])
```

```python

```