Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 13
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 4<sup>th</sup> 2019, 5 pm**

# Problem 1

Implement the `merge` function that will be used in implementing the `merge_sort` function. This function takes two sorted lists of integers, `L1` and `L2`, as inputs and returns a sorted list of the two lists combined.

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
def merge(L1, L2):
    result=[]
    while L1!=[] and L2!=[]:
        if L1[0]<=L2[0]:
            result.append(L1[0])
            del(L1[0])
        else:
            result.append(L2[0])
            del(L2[0])
    result=result+L1+L2
    return result
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert merge([1, 5, 6, 9],[1, 2, 4, 8]) == [1, 1, 2, 4, 5, 6, 8, 9]
assert merge([1,2,3],[]) == [1, 2, 3]
assert merge([],[1,2,3]) == [1, 2, 3]
```

# Problem 2

Implement the `merge_sort` function. It takes a list of integers `L` and returns a sorted list of the same integers. This function recursively divides `L` in half and runs `merge_sort` on each half, then uses `merge` to combine the two sorted halves.

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def merge_sort(L):
    if len(L)==1:
        return L
    else:
        d1=L[:int(len(L)/2)]
        d2=L[int(len(L)/2):]
        return merge(merge_sort(d1),merge_sort(d2))
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 0.5, grade: true, solution: false -->
```python
import random

def generate_list(n, k):
    return [random.randint(1,k) for i in range(n)]

result = %timeit -o merge_sort(generate_list(1000,1000))
assert result.best < 0.01
assert merge_sort([4, 2, 1, 9, 3, 9, 2, 3, 5, 6, 3, 4, 9, 11]) == [1, 2, 2, 3, 3, 3, 4, 4, 5, 6, 9, 9, 9, 11]
```

Output:
```
6.06 ms ± 24.2 µs per loop (mean ± std. dev. of 7 runs, 100 loops each)
```

# Problem 3

Use your implementation of `comb_sort` (as described in class and issued for the most recent assignment) and check, using timing functions, which `shrink factor k` value works most efficiently: 1.1, 1.3, or 1.5?

To select a value simply delete the BEST_K assignment statements you don't want.

**hint**:
<ul>
<li>When you devise test cases be sure that **every** time `comb_sort` is executed it is given a **new** list of randomly ordered integers.
<li>Be sure to use tests of sufficiently large lists (at least 1,000 elements)
</ul>

<!-- grade_id: cell-10e89ce14c311ec1, points: 0, grade: false, solution: true -->
```python
def comb_sort(L, k):
    gap=len(L)-1
    while gap>=1:
        for i in range(0,len(L)-gap):
            if L[i]>L[i+gap]:
                L[i],L[i+gap]=L[i+gap],L[i]
        gap=int(gap/k)
    
    while True:
        break_or_not=1
        for i in range(len(L)-1):
            if L[i]>L[i+1]:
                L[i],L[i+1]=L[i+1],L[i]
                break_or_not=0
        if break_or_not==1:
            break
    return L 

### SELECT YOUR ANSWER BY REMOVING THE OTHER TWO ASSIGNMENTS

BEST_K = 1.3

import random

def generate_list(lenth, maximum):
    return [random.randint(1,maximum) for i in range(lenth)]

k1_1=%timeit comb_sort(generate_list(1000,100),1.1)
k1_3=%timeit comb_sort(generate_list(1000,100),1.3)
k1_5=%timeit comb_sort(generate_list(1000,100),1.5)
```

Output:
```
6.18 ms ± 64.4 µs per loop (mean ± std. dev. of 7 runs, 100 loops each)
4.21 ms ± 44.1 µs per loop (mean ± std. dev. of 7 runs, 100 loops each)
9.26 ms ± 420 µs per loop (mean ± std. dev. of 7 runs, 100 loops each)
```

<!-- grade_id: cell-3eef655e969a4e9c, points: 1, grade: true, solution: false -->
```python
# There is a hidden test here
```