Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 14
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 5<sup>th</sup> 2019, 5 pm**

# Problem 1

Implement functions that will allow you to use the `heap` data structure. A new `heap` begins as a Python list with its 1<sup>st</sup> element set to `None`.

The heap will be handled through two functions:

(1) `insert` - takes a heap `h` and integer `val` as input and returns the heap with the integer inserted into the appropriate location of the heap.

(2) `delete` - takes a heap `h` as input, returns the integer at the root, and deletes that integer from the heap.

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
def insert(h, val):
    h.append(val)
    index=len(h)-1
    if len(h)>2:
        while h[int(index/2)]<h[index]:
            h[int(index/2)],h[index]=h[index],h[int(index/2)]
            index=int(index/2)
            if index==1:
                break


def delete(h):
    root=h[1]
    h[1],h[len(h)-1]=h[len(h)-1],h[1]
    del(h[len(h)-1])
    index=1
    while index<=int(len(h)/2-1):
        if h[2*index]<h[2*index+1]:
            if h[index]<h[2*index+1]:
                h[index],h[2*index+1]=h[2*index+1],h[index]
                index=2*index+1
            else:
                break
        elif h[2*index]>=h[2*index+1]:
            if h[index]<=h[2*index]:
                h[index],h[2*index]=h[2*index],h[index]
                index=2*index
            else:
                break
    if index==int((len(h)-1)/2):
        if h[index]<h[2*index]:
            h[index],h[2*index]=h[2*index],h[index]
    return root
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 1.5, grade: true, solution: false -->
```python
heap = [None]
insert(heap, 9)
insert(heap, 8)
insert(heap, 1)
insert(heap, 3)
insert(heap, 10)
insert(heap, 5)
assert heap == [None, 10, 9, 5, 3, 8, 1]
val = delete(heap)
assert val == 10
assert heap == [None, 9, 8, 5, 3, 1]
val = delete(heap)
assert val == 9
assert heap == [None, 8, 3, 5, 1]
val = delete(heap)
assert val == 8
assert heap == [None, 5, 3, 1]
```

# Problem 2

Implement the `heap_sort` function using the heap object you just defined.

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def heap_sort(L):
    h=[None]
    for val in L:
        insert(h, val)
    result=[]
    while len(h)>=2:
        result.append(delete(h))
    return result
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 0.5, grade: true, solution: false -->
```python
import random

def generate_list(n, k):
    return [random.randint(1,k) for i in range(n)]

L = [4, 2, 1, 6, 8, 9, 2, 4, 5, 2, 3, 3]
result = %timeit -o heap_sort(generate_list(1000,1000))
assert heap_sort(L) == [9, 8, 6, 5, 4, 4, 3, 3, 2, 2, 2, 1]
assert result.best < 0.03
```

Output:
```
11.6 ms ± 67.6 µs per loop (mean ± std. dev. of 7 runs, 100 loops each)
```

```python

```