Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 14
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**April 8<sup>th</sup> 2019**

# Problem 1

Write a function called `is_heap` that takes a list of distinct integers `L` as input and returns `True` if it is an implicit heap and `False` otherwise.

**note**: You can assume that the first element of the list is set to `None`. Actual values, if they exist, populate the list starting from index 1.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def is_heap(L):
    if len(L)%2==0:
        for i in range(1,len(L)//2):
            if L[i]<L[2*i] or L[i]<L[2*i+1]:
                return False
        return True
    else:
        for i in range(1,(len(L))//2-1):
            if L[i]<L[2*i]or L[i]<L[2*i+1]:
                return False
            if L[-1]>L[(len(L)-1)//2]:
                return False                
        return True
```

```python
assert is_heap([None,1,2,3,4,5,6,7]) == False
assert is_heap([None,4,2,6,1,3,5,7]) == False
assert is_heap([None,4,2,6,1,3,5,7,0]) == False
assert is_heap([None,1]) == True
assert is_heap([None]) == True
```

```python
def is_heap(L):                                                    # ????????????????????
    def helper(L,i):
        if i>len(L):
            return True
        s=i*2
        b=i*2+1
        if s<=len(L):
            
        if b<=len(L):
            
    if L[i]>L[s] and L[i]>L[b]:
        return helper(L,s) and helper(L,b) 
```

```python
def is_heap(L):
    def helper(L,i):
        index1=2*i
        index2=2*i+1
        if index1>len(L)-1 or index2>len(L)-1:
            return True
        else:
            if L[index1]>L[i] or L[index2]>L[i]:
                return False
            if L[index1]<=L[i] and L[index2]<=L[i]:
                return helper(L,2)
    return helper(L,1)
```

```python
assert is_heap([None,1,2,3,4,5,6,7]) == False
assert is_heap([None,4,2,6,1,3,5,7]) == False
assert is_heap([None,4,2,6,1,3,5,7,0]) == False
assert is_heap([None,1]) == True
assert is_heap([None]) == True
```

```python
# The following validation is wrong 
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert is_heap([None,1,2,3,4,5,6,7]) == False
assert is_heap([None,4,2,6,1,3,5,7]) == True
assert is_heap([None,4,2,6,1,3,5,7,0]) == True
assert is_heap([None,1]) == True
assert is_heap([None]) == True
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-13-039bf4a2cd6e> in <module>()
      1 assert is_heap([None,1,2,3,4,5,6,7]) == False
----> 2 assert is_heap([None,4,2,6,1,3,5,7]) == True
      3 assert is_heap([None,4,2,6,1,3,5,7,0]) == True
      4 assert is_heap([None,1]) == True
      5 assert is_heap([None]) == True

AssertionError:
```