Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 15
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 9<sup>th</sup> 2019, 5 pm**

# Racquetball Simulation

The goal is to simulate multiple games of racquetball where the ability of the two opponents is represented by the probability that they win a point when they are serving. The inputs to the simulation are the probability for player A, the probability for player B, and the number of games to simulate. The output is a nicely formatted summary of the results.

We end a game when one of the players reaches a total of 15 points. We will also consider shutouts. If one player gets to 7 before the other player has scored a point, the game ends. Our simulation should keep track of both the number of wins for each player and the number of wins that are shutouts.

Please fill in the missing parts of the following code and run the result. An example execution would be similar to this:

<pre>
Prob A is 0.6, Prob B is 0.5
Summary of 500  games:

            wins (% total )      shutouts (% wins ) 
--------------------------------------------
 Player A:  319 (63.8%)          42 (13.2%)
 Player B:  181 (36.2%)           8 (4.4%)
</pre>

# Problem 1

Implement the class `Player` that will be used to represent a racquetball player. This class should keep track of the player's probability of winning when serving (`prob`) and the player's current score (`score`).

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
import random

class Player :
# A Player keeps track of service probability and score
    def __init__(self, prob):
    # Create a player with this probability
        self.prob=prob
        self.score=0
    def winsServe(self):
    # Returns a Boolean that is true with probability self.prob
        compete=random.randint(1,100)
        if 1<=compete<=100*self.prob:
            return True
        if 100*self.prob<compete<=100:
            return False

    def incScore(self):
    # Add a point to this player's score
        self.score=self.score+1
        
    def getScore(self):
    # Returns this player's current score
        return self.score
```

<!-- grade_id: cell-b1b12dbb4976108d, points: 0.5, grade: true, solution: false -->
```python
ming = Player(0.51)
daniel = Player(0.49)
while ming.getScore() < 15 and daniel.getScore() < 15:
    if ming.winsServe():
        ming.incScore()
    else:
        daniel.incScore()
assert ming.prob == 0.51
assert daniel.prob == 0.49
assert max(ming.getScore(),daniel.getScore()) == 15
```

# Problem 2

Implement the class `RBallGame` that will be used to represent a racquetball game. This class should keep track of two players (`playerA` and `playerB`) along with the the player who is **currently** the server (`server`).

<!-- grade_id: cell-26254032e02425bf, points: 0, grade: false, solution: true -->
```python
class RBallGame:
# A RBallGame represents a game in progress. A game has two players
# and keeps track of which one is currently serving.
    def __init__(self, probA, probB):
    # Create a new game having players with the given probs.
    # Player A always serves first
        self.player_A=Player(probA)
        self.player_B=Player(probB)
        self.server=self.player_A 
        
    def play(self):
    # Play the game to completion
        if self.__isOver()==False:                                                                              
            if self.server.winsServe()==True:
                self.server.incScore()
            else:
                self.changeServer()
            return self.play()

    def __isOver(self):
    # Returns game is finished ( i.e. one of the players has won).
        if (self.player_A.getScore()==7 and self.player_B.getScore()==0) or (self.player_B.getScore()==7 and self.player_A.getScore()==0) or (self.player_A.getScore()==15) or (self.player_B.getScore()==15):
            return True
        else:
            return False
    
    def changeServer(self):
    # Switch which player is serving
        if self.server==self.player_A:
            self.server=self.player_B
        else:
            self.server=self.player_A
            
    def getScores(self):
    # Returns the current scores of player A and player B
        return self.player_A.getScore(),self.player_B.getScore()
```

<!-- grade_id: cell-68689597d1269994, points: 1, grade: true, solution: false -->
```python
game = RBallGame(0.51, 0.49)
game.play()
assert max(game.getScores()) in [7,15]
game.changeServer()
game.play()
assert max(game.getScores()) in [7,15]
```

# Problem 3

Implement the class `SimStats` that will be used to represent the results of a series of racquetball games between two players of certain service win probabilities. This class should keep track of the number of wins and shutouts each player has.

<!-- grade_id: cell-cc09a13313390cc2, points: 0, grade: false, solution: true -->
```python
class SimStats:
# SimStats handles the accumulation of statistics across multiple
# (completed) games. It tracks the wins and shutouts for each
# player.
    def __init__(self):
        # Create a new accumulator for a series of games
        self.winsA=0
        self.winsB=0
        self.shutsA=0
        self.shutsB=0
        
    def update(self, aGame):
        # Determine the outcome of a Game and update statistics
        if aGame.getScores()[0]==15:
            self.winsA+=1
        if aGame.getScores()[1]==15:
            self.winsB+=1
        if aGame.getScores()[0]==7 and aGame.getScores()[1]==0:
            self.shutsA+=1
            self.winsA+=1
        if aGame.getScores()[0]==0 and aGame.getScores()[1]==7:
            self.shutsB+=1
            self.winsB+=1
            
    def printReport(self):
    # Print a nicely formatted report
        n = self.winsA + self.winsB
        
        print(" Summary of" , n , " games : \n")
        print("            wins (% total )      shutouts (% wins ) ")
        print("--------------------------------------------")
        self.__printLine ("A" , self.winsA , self.shutsA , n)
        self.__printLine ("B" , self.winsB , self.shutsB , n)

    def __printLine(self , label , wins , shuts , n):
        template = " Player {0}:{1:5} ({2:5.1%}) {3:11} ({4})"
        if wins == 0: # Avoid division by zero!
            shutStr = "-----"
        else:
            shutStr = "{0:4.1%}".format(float(shuts)/wins)
        print (template.format(label , wins , float (wins)/n , shuts , shutStr) )
```

<!-- grade_id: cell-e031a3cd64785fc0, points: 0.5, grade: true, solution: false -->
```python
def printintro():
    print("This program simulates games of racquetball between two")
    print('players called "A" and "B". The ability of each player is')
    print("indicated by a probability (a number between 0 and 1) that")
    print("the player wins the point when serving. Player A always")
    print("has the first serve at the beginning of each game.\n")
    
def main():
    printintro()
    probA = 0.51 # the prob player A wins a serve
    probB = 0.49 # the prob player B wins a serve
    n = 500 # How many games to simulate?

    # Play the games
    print ("ProbA is {}, ProbB is {}".format(probA, probB))
    stats = SimStats()
    for i in range(n):
        theGame = RBallGame(probA, probB)
        theGame.play()
        stats.update(theGame)
    # Print the results
    stats.printReport()

main()
```

Output:
```
This program simulates games of racquetball between two
players called "A" and "B". The ability of each player is
indicated by a probability (a number between 0 and 1) that
the player wins the point when serving. Player A always
has the first serve at the beginning of each game.

ProbA is 0.51, ProbB is 0.49
 Summary of 500  games : 

            wins (% total )      shutouts (% wins ) 
--------------------------------------------
 Player A:  281 (56.2%)          32 (11.4%)
 Player B:  219 (43.8%)           6 (2.7%)
```

```python

```