Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 16
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 10<sup>th</sup> 2019, 5 pm**

# Problem 1

Implement the `IntegerSet` class that can contain integers in the range 0 through 100. Represent the set internally as a list `set` of `boolean` values. Element `set[i]` is `True` if integer `i` is in the set. Element `set[i]` is `False` if integer `i` is not in the set. The default constructor initializes the `set` to the so-called 'empty set,' i.e., a set for which all the elements are `False`.

This class should also have the following methods:

`unionOfSets`: updates the current set with the set-theoretic union of itself and another set (i.e., an element of the result is set to `True` if that element is `True` in either or both of the existing sets, and an element of the result is set to `False` if that element is `False` in both of the existing sets).

`intersectionOfSets`: updates the current set with the set-theoretic intersection of itself and another set (i.e., an element of the result is set to `False` if that element is `False` in either or both of the existing sets, and an element of the result is set to `True` if that element is `True` in each of the existing sets).

`insertElement`: inserts an integer `k` by setting `set[k]` to `True`.

`deleteElement`: deletes an integer `k` by setting `set[k]` to `False`.

`__str__`: formats the set as a list of numbers separated by spaces. If the set is empty then return `-`.

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
class IntegerSet1:
    
    def __init__(self):
        self.set = [False for i in range(101)]
        
    def unionOfSets(self,other):
        for i in range(101):
            if other.set[i]==True and self.set[i]==False:
                self.set[i]=True
                
    def intersectionOfSets(self,other):
        for i in range(101):
            if other.set[i]==False and self.set[i]==True:
                self.set[i]=False
                
    def insertElement(self,k):
        self.set[k]=True
    
    def deleteElement(self,k):
        self.set[k]=False
        
    def __str__(self):
        if True not in self.set:
            return 'IntegerSet: -'
        else:
            result=''
            for i in range(101):
                if self.set[i]==True:
                    result=result+' '+str(i)
            return 'IntegerSet:'+ result
    
    def __getitem__(self,i):
        return self.set[i]
```

<!-- grade_id: cell-b1b12dbb4976108d, points: 1, grade: true, solution: false -->
```python
a = IntegerSet1()
a.set == [False for i in range(101)]
assert str(a) == 'IntegerSet: -'

for i in range(101):
    assert a[i] == False

for i in range(0, 101, 2):
    a.insertElement(i)
    
for i in range(101):
    assert a[i] == (i%2==0)

assert str(a) == 'IntegerSet: 0 2 4 6 8 10 12 14 16 18 20 22 24 26 28 30 32 34 36 38 40 42 44 46 48 50 52 54 56 58 60 62 64 66 68 70 72 74 76 78 80 82 84 86 88 90 92 94 96 98 100'

for i in range(50, 101):
    a.deleteElement(i)

assert str(a) == 'IntegerSet: 0 2 4 6 8 10 12 14 16 18 20 22 24 26 28 30 32 34 36 38 40 42 44 46 48'

b = IntegerSet1()
for i in range(0, 50, 3):
    b.insertElement(i)

assert str(b) == 'IntegerSet: 0 3 6 9 12 15 18 21 24 27 30 33 36 39 42 45 48'

a.unionOfSets(b)
assert str(a) == 'IntegerSet: 0 2 3 4 6 8 9 10 12 14 15 16 18 20 21 22 24 26 27 28 30 32 33 34 36 38 39 40 42 44 45 46 48'

a.intersectionOfSets(b)
assert str(a) == 'IntegerSet: 0 3 6 9 12 15 18 21 24 27 30 33 36 39 42 45 48'
```

# Problem 2

Implement the `IntegerSet` class with a list of integers.

<!-- grade_id: cell-26254032e02425bf, points: 0, grade: false, solution: true -->
```python
class IntegerSet2:
    def __init__(self):
        self.set = []
    
    def unionOfSets(self,other):
        result=[]
        for i in range(101):
            if i in self.set or i in other.set:
                result.append(i)
        a.set=result
                
    def intersectionOfSets(self,other):
        Result=[]
        for i in range(101):
            if i in self.set and i in other.set:
                Result.append(i)
        a.set=Result
                
    def insertElement(self,k):
        self.set.append(k)
    
    def deleteElement(self,k):
        if k in self.set:
            self.set.remove(k)
        
    def __str__(self):
        if len(self.set)==0:
            return 'IntegerSet: -'
        else:
            result=''
            for num in self.set:
                result=result+' '+str(num)
            return 'IntegerSet:'+ result 
        
    def __getitem__(self,i):
        if i in self.set:
            return True
        else:
            return False
```

<!-- grade_id: cell-68689597d1269994, points: 1, grade: true, solution: false -->
```python
a = IntegerSet2()
a.set == []
assert str(a) == 'IntegerSet: -'

for i in range(101):
    assert a[i] == False

for i in range(0, 101, 2):
    a.insertElement(i)
    
for i in range(101):
    assert a[i] == (i%2==0)

assert str(a) == 'IntegerSet: 0 2 4 6 8 10 12 14 16 18 20 22 24 26 28 30 32 34 36 38 40 42 44 46 48 50 52 54 56 58 60 62 64 66 68 70 72 74 76 78 80 82 84 86 88 90 92 94 96 98 100'

for i in range(50, 101):
    a.deleteElement(i)

assert str(a) == 'IntegerSet: 0 2 4 6 8 10 12 14 16 18 20 22 24 26 28 30 32 34 36 38 40 42 44 46 48'

b = IntegerSet2()
for i in range(0, 50, 3):
    b.insertElement(i)

assert str(b) == 'IntegerSet: 0 3 6 9 12 15 18 21 24 27 30 33 36 39 42 45 48'

a.unionOfSets(b)
assert str(a) == 'IntegerSet: 0 2 3 4 6 8 9 10 12 14 15 16 18 20 21 22 24 26 27 28 30 32 33 34 36 38 39 40 42 44 45 46 48'

a.intersectionOfSets(b)
assert str(a) == 'IntegerSet: 0 3 6 9 12 15 18 21 24 27 30 33 36 39 42 45 48'
```

```python

```