# Homework 17
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 11<sup>th</sup> 2019, 5 pm**

# Problem 1

Implement the `stack` class using linked lists. This will be more efficient for pushing and popping than a plain Python list. You should implement the following methods:

`__init__`: create a new `stack` with no elements.

`pop`: return the element that was inputted last and remove the element from the `stack`.

`peek`: return the element that was inputted last.

`push`: insert a new element at the end of the `stack`.

`__str__`: returns a string representation of the `stack`.

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
class node(object):
    def __init__(self,value):
        self.value=value
        self.next=None

class stack():
    
    def __init__(self):
        self.head=None
        self.tail=None
    
    def pop(self):
        if self.head==self.tail and self.head!=None: # len==1
            store=self.tail
            self.tail=None
            self.head=None
            return store.value
        
        if self.head!=self.tail:    # len>1
            current_N=self.head
            previous_N=None
            store=self.tail
            while current_N.next!=None:  # current_N 在 tail的位置
                previous_N=current_N
                current_N=current_N.next
            previous_N.next=None
            self.tail=previous_N
            return store.value
    
    def peek(self):
        if self.head!=None:     # len!=0
            return self.tail.value
        
    def push(self,value):
        if self.tail==None:     # len==0
            new_N=node(value)   # create one variable instead of calling it twice!!!
            self.head=new_N
            self.tail=new_N
        else:
            new_N=node(value)
            self.tail.next=new_N
            self.tail=new_N
            
    def __str__(self):
        current_N=self.head
        result=''
        while current_N!=None:
            result=str(current_N.value)+' '+result
            current_N = current_N.next
        return 'stack: '+result.strip()
```

<!-- grade_id: cell-b1b12dbb4976108d, points: 1, grade: true, solution: false -->
```python
s = stack()
s.push(1)
s.push(2)
s.push(3)
s.push(4)
s.push(5)

assert s.peek() == 5
assert s.pop() == 5
assert s.peek() == 4
assert str(s) == 'stack: 4 3 2 1'

s.pop()
s.pop()
s.pop()
assert str(s) == 'stack: 1'
```

# Problem 2

Implement the `queue` class using linked lists. This will be more efficient for pushing and popping than a plain Python list. You should implement the following methods:

`__init__`:create a new `queue` with no elements.

`pop`: return the element that was inputted first and remove the element from the `queue`.

`peek`: return the element that was inputted first.

`push`: insert a new element at the end of the `queue`.

`__str__`: returns a string representation of the `queue`.

<!-- grade_id: cell-26254032e02425bf, points: 0, grade: false, solution: true -->
```python
class node(object):
    def __init__(self,value):
        self.value=value
        self.next=None

class queue(object):
    def __init__(self):
        self.head=None
        self.tail=None
    
    def pop(self):
        if self.head==self.tail and self.head!=None: # len==1
            store=self.head
            self.tail=None
            self.head=None
            return store.value
        
        if self.head!=self.tail:    # len>1
            store=self.head
            self.head=self.head.next
            return store.value
    
    def peek(self):
        if self.head!=None:     # len!=0
            return self.head.value
        
    def push(self,value):
        if self.tail==None:     # len==0
            new_N=node(value)
            self.head=new_N
            self.tail=new_N
        else:
            new_N=node(value)
            self.tail.next=new_N
            self.tail=new_N
            
    def __str__(self):
        current_N=self.head
        result=''
        while current_N!=None:
            result=result+' '+str(current_N.value)
            current_N = current_N.next
        return 'queue: '+result.strip()
```

<!-- grade_id: cell-68689597d1269994, points: 1, grade: true, solution: false -->
```python
q = queue()
q.push(1)
q.push(2)
q.push(3)
q.push(4)
q.push(5)

assert q.peek() == 1
assert q.pop() == 1
assert q.peek() == 2
assert str(q) == 'queue: 2 3 4 5'

q.pop()
q.pop()
q.pop()
assert str(q) == 'queue: 5'
```

```python

```