Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 19
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 16<sup>th</sup> 2019, 5 pm**

# Problem 1

Given an `DirectedGraphNode`, we can serialize the graph into a string.

We use `#` as a separator for each node and `,` as a separator for each neighbor of the node.
 
As an example, consider the serialized graph {0,1,2#1,2#2,2#}.

The graph has a total of three nodes, and therefore contains three parts separated by `#`.

The first node is labeled 0 and it has two neighbors: node 1 and node 2.
The second node is labeled 1 and it has one neighbor: node 2.
The third node is labeled 2 and it has one neighbor: node 2 (itself).

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
class DirectedGraphNode:
#     all_nodes=[]
    
    def __init__(self,val):
        self.val=val
        self.neighbors=[]
    
#     def append(self,something):
#         self.append(something)
#         self.__class__.all_nodes.append(something)
        
#     def get_nodes(self):
#         result=[self.val]
#         if self.neighbors==[]:
#             return result
#         else:
#             for item in self.neighbors:
#                 result=result+item.get_nodes()
#             return result    
        
#     def print_subgraph_from_here(self,nodelist):
#         result=''
#         nodes=list(set(self.__class__.all_nodes))
#         for item in nodes:
#             copy_n=item.neighbors[:]
#             for i in range(len(copy_n)):
#                 copy_n[i]=str(copy_n[i])
#             result=result+str(item.val)+','+','.join(copy_n)+'#'
#         return result
    
    def print_subgraph_from_here(self,nodelist):
        nodelist.append(self)
        index=0
        while index<len(nodelist):
            for item in nodelist[index].neighbors:
                if item not in nodelist:
                    nodelist.append(item)
            index=index+1
            
        result=''
        for node in nodelist:
            copy_nb=node.neighbors[:]
            for i in range(len(copy_nb)):
                copy_nb[i]=str(copy_nb[i].val)
            result=result+str(node.val)+','+','.join(copy_nb)+'#'
        return result
        
```

<!-- grade_id: cell-b1b12dbb4976108d, points: 1, grade: true, solution: false -->
```python
n0 = DirectedGraphNode(0)
n1 = DirectedGraphNode(1)
n2 = DirectedGraphNode(2)

n0.neighbors.append(n1)
n0.neighbors.append(n2)
n1.neighbors.append(n2)
n2.neighbors.append(n2)

assert n0.print_subgraph_from_here([])=='0,1,2#1,2#2,2#'
```

# Problem 2

Write a function called `clone_graph` that takes the head node of a graph as input and returns a deep copy (clone) of the graph. Each node in the graph contains a label (int) and a list (List[DirectedGraphNode]) of its neighbors. There is an edge between the given node and each of the nodes in neighbors.

<!-- grade_id: cell-26254032e02425bf, points: 0, grade: false, solution: true -->
```python
import copy
def cloneGraph(node):
    return copy.deepcopy(node)
```

```python
def cloneGraph(node):
    copy=DirectedGraphNode(node.val)
    copy.neighbors=node.neighbors[:]   
    return copy
```

```python
def cloneGraph(node):
    copy  = {}
    def DFS(n):
        if n in copy:
            return copy[n]
        node_new = DirectedGraphNode(n.val)
        copy[n] = node_new
        for i in n.neighbors:
            node_new.neighbors.append(DFS(i))
        return node_new
    return DFS(node)
```

<!-- grade_id: cell-68689597d1269994, points: 1, grade: true, solution: false -->
```python
n0 = DirectedGraphNode(0)
n1 = DirectedGraphNode(1)
n2 = DirectedGraphNode(2)
n3 = DirectedGraphNode(3)

n0.neighbors.append(n1)
n0.neighbors.append(n2)
n1.neighbors.append(n2)
n2.neighbors.append(n3)
n3.neighbors.append(n3)

n0_clone = cloneGraph(n0)

assert n0_clone.print_subgraph_from_here([]) == '0,1,2#1,2#2,3#3,3#'
assert n0_clone.print_subgraph_from_here([]) == n0.print_subgraph_from_here([])
assert n0_clone is not n0
```

```python

```