Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 21
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 18<sup>th</sup> 2019, 5 pm**

```python
class Digraph(object):
    # nodes is a list of strings in the graph
    # edges is a dictionary mapping each node to a list of its children
    def __init__(self):
        self.nodes = []
        self.edges = {}

    def childrenOf(self, node):
        return self.edges[node]
    
    def hasNode(self, node):
        return node in self.nodes
    
    def __str__(self):
        result = ''
        for src in self.nodes:
            for dest in self.edges[src]:
                result = result + src.getName() + '->'\
                         + dest.getName() + '\n'
        return result[:-1] #omit final newline
```

# Problem 1

Implement `Dijkstra's Algorithm`.

Given a graph `G`, node `s` and node `e`, return the shortest distance between these two nodes. If there is no path between s and e, return 9999.

<img src="graph.png" width="350">

<!-- grade_id: cell-761dacc86a1b5d0b, points: 0, grade: false, solution: true -->
```python
def dijkstra(G, s, e):
    '''
    G: Graph
    s: start node
    e: end node
    '''
    # table initial setting
    if G.edges[s]==[]:
        return 9999
    else:
        unvisited=[]
        visited=[s]
        for item in G.edges.keys():
            if G.edges[item]!=[]:
                unvisited.append(item)
        unvisited.remove(s)
        table={}
        neighbors=[]
        for tuples in G.edges[s]:
            neighbors.append(tuples[0])
            table[tuples[0]]=tuples[1]
        for node in G.nodes:
            if node not in neighbors:
                table[node]=9999
        table[s]=0
        # change table
        while len(unvisited)!=0:
            # find smallest
            smallest=None
            for item in unvisited:
                if smallest==None or table[item]<smallest[1]:
                    smallest=(item,table[item])
            unvisited.remove(smallest[0])
            visited.append(smallest[0])

            # visit this node and update
            for nb in G.edges[smallest[0]]:
                if nb[0] not in visited:
                    if nb[1]+smallest[1]<table[nb[0]]:
                        table[nb[0]]=nb[1]+smallest[1]

        return table[e]
```

```python
def dijkstra(G, s, e):
    # initialize
    table={item:9999 for item in G.nodes}
    table[s]=0
    unvisited=[node for node in G.nodes]
    # algorithm
    while unvisited:
        smallest=min([(i,table[i]) for i in unvisited],key=lambda x:x[1])[0]
        unvisited.remove(smallest)
        if smallest in G.edges:
            for node, weight in G.edges[smallest]:
                table[node]=min(table[node],table[smallest]+weight)
    return table[e]

```

<!-- grade_id: cell-b1b12dbb4976108d, points: 1, grade: true, solution: false -->
```python
G1 = Digraph()
G1.nodes = ['A', 'B', 'C', 'D']
G1.edges = {
    'A':[('B', 1), ('C', 4), ('D', 6)],
    'B':[('C', 2), ('D', 7)],
    'C':[('D', 1)],
}

assert dijkstra(G1, 'A', 'C') == 3
assert dijkstra(G1, 'B', 'D') == 3
```

<!-- grade_id: cell-22bba7a04540e005, points: 1, grade: true, solution: false -->
```python
G2 = Digraph()
G2.nodes = ['A', 'B', 'C', 'D', 'E', 'F']
G2.edges = {
    'A':[('B', 7), ('C', 9), ('F', 14)],
    'B':[('C', 10), ('D', 15)],
    'C':[('D', 11), ('F', 2)],
    'D':[('E', 6)],
    'E':[],
    'F':[('E', 9)],
}

assert dijkstra(G2, 'A', 'D') == 20
assert dijkstra(G2, 'A', 'E') == 20
assert dijkstra(G2, 'B', 'D') == 15
assert dijkstra(G2, 'B', 'A') == 9999
assert dijkstra(G2, 'C', 'E') == 11
assert dijkstra(G2, 'E', 'A') == 9999
assert dijkstra(G2, 'F', 'E') == 9
```

```python

```