Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 11
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 2<sup>nd</sup> 2019, 5 pm**

# Problem 1

Write a function called `even_odd` that takes a list of integers `L` and returns the **same** list with all the even numbers in front and all the odd numbers in back.

Implement the following algorithm:

<ol>
    <li>Starting from the beginning of `L` check each number
    <li>If the number is even then swap it with each previous number until it reaches the beginning of `L`
    <li>Repeat this process until the all numbers in `L` are checked
</ol>

Here is an example:

<pre>
even_odd([1, 2, 3, 4])
</pre>

This will produce the following changes at each iteration of the algorithm:

<pre>
[1, 2, 3, 4]   - 1st iteration: nothing changes because 1 is odd
[2, 1, 3, 4]   - 2nd iteration: 2 is even, so it is sent to the front
[2, 1, 3, 4]   - 3rd iteration: nothing changes because 3 is odd
[4, 2, 1, 3]   - 4th iteration: 4 is even, so it is sent to the front
</pre>

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
def even_odd(L):
    for i in range(len(L)):
        if L[i]%2==0:
            while i>0:
                L[i],L[i-1]=L[i-1],L[i]
                i=i-1
    return L
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
import random

def gen_list(size):
    return [random.randint(1,100) for i in range(size)]

assert even_odd([1,2,3,4,5,6]) == [6,4,2,1,3,5]
assert even_odd([1,3,5,7,9,11]) == [1,3,5,7,9,11]
assert even_odd([2,4,6,8,10,12]) == [12,10,8,6,4,2]

result = %timeit -o even_odd(gen_list(2500))
assert result.best < 1
```

Output:
```
357 ms ± 5.67 ms per loop (mean ± std. dev. of 7 runs, 1 loop each)
```

# Problem 2

What is the time complexity of `even_odd`?

**note**: please delete all incorrect options below so that you have a single line that represents your final answer

<!-- grade_id: unique, points: 0, grade: false, solution: true -->
```python
EVEN_ODD_COMPLEXITY = 'polynomial'
```

<!-- grade_id: unique1, points: 0.5, grade: true, solution: false -->
```python
# There is a hidden test here
```

# Problem 3

Can you write a more efficient version of `even_odd`?

Write a function called `even_odd_fast` that performs the same task as `even_odd` but runs more efficiently. This function starts from the beginning and end of the list and loops while **beg is less than end**.

<pre>
[ 1, 3, 2, 4 ]
  ^        ^
 beg      end
</pre>

<ol>
    <li>Check if the integer at beg is even. If it is even we can assume that it is in the right location and move beg forward.
    <li>Check if the integer at end is odd. If it is odd we can assume that it is in the right location and move end backward.
    <li>If the integer at beg is odd and the integer at end is even then these integers should be **swapped**. Then beg can move forward and end can move backward.
</ol>

Here is what happens at each iteration of the loop:

<pre>
[ 1, 3, 2, 4 ]  - beg and end should be swapped
  ^        ^
 beg      end
 
[ 4, 3, 2, 1 ]  - beg should move forward and end should move backward
  ^        ^
 beg      end
 
[ 4, 3, 2, 4 ]  - beg and end should be swapped
     ^  ^
   beg  end
   
[ 4, 2, 3, 1 ]  - beg should move forward and end should move backward
     ^  ^
   end  beg

** NOW the loop finishes because beg is no longer before end.
</pre>

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def even_odd_fast(L):
    beg=0
    end=-1
    while len(L)+end>beg:
        if L[beg]%2==0:
            beg=beg+1
        if L[end]%2!=0:
            end=end-1
        if len(L)+end>beg:
            if L[beg]%2!=0 and L[end]%2==0:
                L[beg],L[end]=L[end],L[beg]
                beg=beg+1
                end=end-1
    return L
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 0.5, grade: true, solution: false -->
```python
assert even_odd_fast([1,2,3,4,5,6,7,8,9]) == [8,2,6,4,5,3,7,1,9]
assert even_odd_fast([2,4,6,8,10,12,14]) == [2,4,6,8,10,12,14]
assert even_odd_fast([1,3,5,7,9,11,13]) == [1,3,5,7,9,11,13]

result = %timeit -o even_odd_fast(gen_list(2500))
assert result.best < .05
```

Output:
```
3.98 ms ± 43.1 µs per loop (mean ± std. dev. of 7 runs, 100 loops each)
```

# Problem 4

What is the time complexity of `even_odd_fast`?

**note**: please delete all incorrect options below so that you have a single line that represents your final answer

<!-- grade_id: unique11, points: 0, grade: false, solution: true -->
```python
EVEN_ODD_FAST_COMPLEXITY = 'linear'
```

<!-- grade_id: unique12, points: 0.5, grade: true, solution: false -->
```python
# There is a hidden test here
```