Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 9
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **March 27<sup>th</sup> 2019, 5 pm**

# TIC-TAC-TOE

**Tic-Tac-Toe** is a game for two players who take turns marking spaces in a $3x3$ grid. One player uses 'X' to mark spaces and the other player uses 'O' to mark spaces. The player who succeeds in placing three of their marks in a horizontal, vertical, or diagonal row wins the game. If all the spaces are marked and there is no winner, the game is a draw.

In this homework assignment you will implement three functions that are self-contained and can be used as parts of an overall program that let's one play Tic-Tac-Toe.

Wikipedia: https://en.wikipedia.org/wiki/Tic-tac-toe

# Problem 1: Display Grid

Write a function called `display_grid` that takes a one-dimensional list `L` that contain 'X's, 'O's, or ' 's, and an integer `r` that represents the number of rows. This function will return a string with proper formatting escape characters that can be printed as a 2-dimensional matrix with `r` rows.

** You can assume that the length of `L` is evenly divisible by `r`. So when `L` has 12 elements and `r` is 4, `display_grid` will return a string that, when printed, will display a 4 x 3 matrix.

<pre>
display_grid(['X','O','X',' ',' ','X','O','O',' '],3)

'-----------------\n| X | O | X |   |\n-----------------\n|   | X | O | O |\n-----------------\n|   |   |   | X |\n-----------------'

print(display_grid(['X','O','X',' ',' ','X','O','O',' '],3))

-------------
| X | O | X |
-------------
|   |   | X |
-------------
| O | O |   |
-------------
</pre>

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
def display_grid(L, r):
    each_row=int(len(L)/r)
    length_row=4*each_row+1
    result=['-'*length_row]
    body=[]
    for item in L:
        body.append('| '+item+' |')
    modified_body=[]
    if body !=[]:
        for i in range(r):
            modified_body.append(body[0:each_row])
            del(body[0:each_row])
# row=[]
# all_row=[]
# for i in L:
#     row.append(L[i])
#     if (i+1)%each_row==0:
#         all_rows.append(row)
#         row=[]

    MB=[]
    for item in modified_body:
        MB.append(''.join(item))
    new_MB=[]
    for string in MB:
        new_MB.append(string.replace('||','|'))
    for item in new_MB:
        result.append('\n'+item+'\n'+'-'*length_row)
            
    return ''.join(result)
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 1, grade: true, solution: false -->
```python
assert display_grid(['X','O','X',' ',' ','X','O','O',' '],3) == '-------------\n| X | O | X |\n-------------\n|   |   | X |\n-------------\n| O | O |   |\n-------------'
assert display_grid(['X','O','X',' ',' ','X','O','O',' ','O','X',' '],3) == '-----------------\n| X | O | X |   |\n-----------------\n|   | X | O | O |\n-----------------\n|   | O | X |   |\n-----------------'
assert display_grid(['X','O','X',' ',' ','X','O','O',' ','O','X',' '],4) == '-------------\n| X | O | X |\n-------------\n|   |   | X |\n-------------\n| O | O |   |\n-------------\n| O | X |   |\n-------------'
```

# Problem 2: Check for Winner

Write a function called `is_winner` that takes a one-dimensional list `L` that contains 'X's, 'O's, or ' 's, and a 1-character string `player` that is either an 'X' or an 'O'. This function will return a boolean value - `True` if `player` won and `False` otherwise.

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def is_winner(L, player):
    for i in range(0,3):
        if L[i]==L[i+3]==L[i+6]==player:
            return True
    if L[0]==L[4]==L[8]==player:
        return True
    if L[2]==L[4]==L[6]==player:
        return True
    return False
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 0.5, grade: true, solution: false -->
```python
assert is_winner(['X','O','X',' ',' ','X','O','O',' '], 'X') == False
assert is_winner(['X','O','X',' ',' ','X','O','O','X'], 'X') == True
assert is_winner(['X','O','X',' ',' ','X','O','O','X'], 'O') == False
assert is_winner(['X','O','O','X','O','X','O','O','X'], 'O') == True
assert is_winner(['X','O','O','X','O','X','X','O','X'], 'X') == True
assert is_winner(['X','O','O','X','O','X','X','O','X'], 'O') == True
```

# Problem 3: Possible Moves

Write a function called `possible_moves` that takes a one-dimensional list `L` that contains 'X's, 'O's, or ' 's. This function returns a list of empty positions in the grid.

** Positions correspond to the numbers in the following grid:

<pre>
-------------
| 0 | 1 | 2 |
-------------
| 3 | 4 | 5 |
-------------
| 6 | 7 | 8 |
-------------
</pre>

<!-- grade_id: cell-605908b0deafec1e, points: 0, grade: false, solution: true -->
```python
def possible_moves(L):
    result=[]
    for i in range(len(L)):
        if L[i]==' ':
            result.append(i)
    return result
```

<!-- grade_id: cell-dda05966fb20db3c, points: 0.5, grade: true, solution: false -->
```python
assert possible_moves(['X','O','X',' ',' ','X','O','O','X']) == [3, 4]
assert possible_moves([' ',' ',' ',' ',' ','X','O',' ',' ']) == [0, 1, 2, 3, 4, 7, 8]
assert possible_moves([' ',' ',' ',' ',' ',' ',' ',' ',' ']) == [0, 1, 2, 3, 4, 5, 6, 7, 8]
assert possible_moves(['X','O','X','O','X','X','O','O','X']) == []
```