Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 12
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 3<sup>rd</sup> 2019, 5 pm**

# Problem 1

Implement the `insert` function as it works in the insertion sort algorithm. This function takes three inputs: (i) a list of integers `L`, (ii) an integer `k`, and (iii) an integer `i`.

`L` is divided into two sections. The first section, marked by the range 0 to `k`, represents the sorted section. The second section, marked by the range `k`+1 to -1, represents the unsorted section.

`k` represents the index of the last element in the sorted section of `L`.

`i` represents the index of the integer that is to be inserted into the sorted section of `L`.

This function returns `L` with the integer at `i` inserted into the sorted section.

** Do <strong>NOT</strong> use the built in Python functions: `insert`, `del`, `pop`, or `remove`!

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
# special case

# def insert(L, k, i)
#     while L[i]<L[i-1]:
#         L[i],L[i-1]=L[i-1],L[i]
#         i=i-1 
#         if i==0:
#             break
#     return L
```

```python
# general case

def insert(L, k, i):
    original_k=k
    L[i],L[k+1]=L[k+1],L[i]
    while L[k+1]<L[k]:
        L[k+1],L[k]=L[k],L[k+1]
        k=k-1 
        if k==-1:
            break
    k=original_k+1
    return L
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
L = [1, 3, 7, 2, 7, 9, 5, 4, 8]
insert(L, 2, 3)
assert L == [1, 2, 3, 7, 7, 9, 5, 4, 8]

L = [1, 2, 3, 7, 7, 9, 5, 4, 8]
insert(L, 5, 6)
assert L == [1, 2, 3, 5, 7, 7, 9, 4, 8]

L = [9, 2, 3, 7, 6, 9, 5, 4, 8]
insert(L, 0, 1)
assert L == [2, 9, 3, 7, 6, 9, 5, 4, 8]
```

# Problem 2

Implement the `insertion_sort` function. This function takes a list of integers `L` as input and updates `L` so that it is sorted. This function must use the `insert` function.

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def insertion_sort(L):
    for i in range(1,len(L)):
        insert(L, i-1, i)
    return L
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 0.5, grade: true, solution: false -->
```python
L = [6, 3, 9, 5, 7, 0, 9, 7, 5, 4, 3, 1, 3, 4, 2, 4, 6, 1, 7]
insertion_sort(L)
assert L == [0, 1, 1, 2, 3, 3, 3, 4, 4, 4, 5, 5, 6, 6, 7, 7, 7, 9, 9]

L = [0, 0, 0, 0, 0, 0, -1, 1, 1, 1, 1, 1, 1, -1, -1, -1, -1, -1, -1]
insertion_sort(L)
assert L == [-1, -1, -1, -1, -1, -1, -1, 0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1]
```

# Problem 3

Implement the `comb_sort` function. In `bubble sort`, when any two elements are compared, they always have a gap (distance from each other) of 1. The basic idea of comb sort is that the gap can be **much more than 1**. The inner loop of bubble sort, which does the actual swap, is modified such that gap between swapped elements goes down (for each iteration of outer loop) in steps of a "shrink factor" `k`.

The gap starts out as the length of the list being sorted and one pass of the aforementioned modified bubble sort is applied with that gap. Then the gap is divided by the shrink factor `k`, the list is sorted with this new gap, and the process repeats until the gap is 1. At this point, comb sort **continues using a gap of 1** until the list is fully sorted. The final stage of the sort is thus equivalent to a bubble sort, but by this time most turtles have been dealt with, so a bubble sort will be more efficient.

<!-- grade_id: cell-10e89ce14c311ec1, points: 0, grade: false, solution: true -->
```python
def comb_sort(L, k):
    gap=len(L)-1
    while gap>=1:
        for i in range(0,len(L)-gap):
            if L[i]>L[i+gap]:
                L[i],L[i+gap]=L[i+gap],L[i]
        gap=int(gap/k)
    
    while True:
        break_or_not=1
        for i in range(len(L)-1):
            if L[i]>L[i+1]:
                L[i],L[i+1]=L[i+1],L[i]
                break_or_not=0
        if break_or_not==1:
            break
    return L       
```

<!-- grade_id: cell-10e89ce14c311ec1, points: 0, grade: false, solution: true -->
```python
def comb_sort(L, k):
    gap=len(L)-1
    while gap>=1:
        for i in range(0,len(L)-gap):
            if L[i]>L[i+gap]:
                L[i],L[i+gap]=L[i+gap],L[i]
        gap=int(gap/k)
    
    for i in range(len(L)-1):
        if L[i]>L[i+1]:
            L[i],L[i+1]=L[i+1],L[i]

    return L       
```

```python
# CR's     (limitation: if k is very big, cannot ensure the examination process when gap=1)

def comb_sort(L, k):
    l=len(L)-1
    while l>=1:
        i=0
        while (i+l)<len(L):
            if L[i]>L[i+l]:
                L[i],L[i+l]=L[i+l],L[i]
            i+=1
        l=int(l//k)
    return L
```

<!-- grade_id: cell-3eef655e969a4e9c, points: 1, grade: true, solution: false -->
```python
L = [3, 2, 7, 9, 6, 4, 2, 3, 4, 5, 6, 4]
comb_sort(L,1.5)
assert L == [2, 2, 3, 3, 4, 4, 4, 5, 6, 6, 7, 9]

L = [-1, 0, 10, 10, 10, 10, 0, 0, 0, 10, 0]
comb_sort(L,1.5)
assert L == [-1, 0, 0, 0, 0, 0, 10, 10, 10, 10, 10]
```

```python

```