Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = ""
```

---

# Homework 4
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **March 19<sup>th</sup> 2019, 5 pm**

# Preparation

To solve today's homework, you need to get a dataset by downloading the `ca-GrQc.txt.gz` file that can be found here - https://snap.stanford.edu/data/ca-GrQc.txt.gz. Extract it the text file and upload it to the folder of this ipynb file. The name of the extraced file is `CA-GrQc.txt`.

The file CA-GrQc.txt is from the Stanford Large Network Dataset Collection and is a record of co-authorship links for articles in the ArXiv category General Relativity.

Each number represents an author. 

<pre>
FromNodeId	ToNodeId
3466        937
</pre>

This represents the fact that author 3466 co-authored a paepr with author 937.

# Problem 1: All coauthorships

Write a function that creates a list of all coauthorships as a list of lists.

<pre>
[[3466, 937], [3466, 5233], ...]. 
</pre>

#### Hints

Step 1: Create a function that converts a string into a list of two authors, like: <pre>[3466, 937]</pre>

Step 2: Use a for loop to iterate over each line in the file.

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
def get_coauthors():
    file=open('CA-GrQc.txt','r')
    result=[]
    for line in file:
        if line[0]!='#':
            num=line.split('	')
            A=int(num[0])
            B=int(num[1])
            L=[A,B]
            result.append(L)
    return result
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 1, grade: true, solution: false -->
```python
assert get_coauthors()[:10] == [[3466, 937], [3466, 5233], [3466, 8579], [3466, 10310], [3466, 15931], [3466, 17038], [3466, 18720], [3466, 19607], [10310, 1854], [10310, 3466]]
```

# Problem 2: Unique authors

Write a function that uses the `CA-GrQc.txt` to create a list of all __unique__ authors. The list should be sorted.

#### Hints

Note that if the edge 3466–937 is in the data, then the edge 937–3466 will also there. This means that for this task you don't need to consider the second author in the line. You can get all authors simply by collecting the first author in each line, that is the numbers in the `FromNodeId` column.

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def get_authors(L):
    unique_authors=[]
    for pair in L:
        if pair[0] not in unique_authors:
            unique_authors.append(pair[0])
    result=sorted(unique_authors)
    return result
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 1, grade: true, solution: false -->
```python
list_unique = get_authors(get_coauthors())

assert list_unique[:10] == [13, 14, 22, 24, 25, 26, 27, 28, 29, 45]
assert list_unique[59] == 231
assert len(list_unique) == 5242
```