Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 10
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **March 28<sup>th</sup> 2019, 5 pm**

# Problem 1 - Check Player Move

Write a function called `check_move`. This will ensure that the move is: (i) an integer, (ii) a valid move in tic-tac-toe (0-8), and (iii) a move in an unoccupied space.

If the move is not an integer an Exception will be raised with the text: 'Must be an integer'

If the move is not between 0-8 an Exception will be raised with the text: 'Must be between 0-8'

If the move is not in an unoccupied space an Exception will be raised with the text: 'Must be in an unoccupied space: X is taken' where X is the integer representing the occupying space.

Inputs: (i) list `L` of 'X's, 'O'x, and ' 's and (ii) an integer `move`

Output: None

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
def possible_moves(board):
    return [key for key, val in enumerate(board) if val == ' ']

def check_move(board, move):

    try:
        move=int(move)
        board[int(move)]
        if board[int(move)] != ' ':
            raise Exception('Must be in an unoccupied space: ' + str(move) + ' is taken')
    except ValueError:
        raise Exception('Must be an integer')
    except IndexError:
        raise Exception('Must be between 0-8')
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
try:
    check_move(['X','O','X',' ',' ','O',' ','X','O'], 'a')
except Exception as msg:
    assert str(msg) == 'Must be an integer'
try:
    check_move(['X','O','X',' ',' ','O',' ','X','O'], '9')
except Exception as msg:
    assert str(msg) == 'Must be between 0-8'
try:
    check_move(['X','O','X',' ',' ','O',' ','X','O'], '2')
except Exception as msg:
    assert str(msg) == 'Must be in an unoccupied space: 2 is taken'
```

# Problem 2 - Computer Move

Write a function called `AI_move`. This will generate a move based on the following strategy:

1. If there is a move that will allow the computer to win, return this move.

2. If there is a move that will allow the human to win in the next turn, return this move.

3. If there is an unoccupied corner (0, 2, 6, 8) return a random corner.

4. If the center (4) is unoccupied return the center.

5. Otherwise, return a random unoccupied space.

You can assume that the computer plays 'O' and the human plays 'X'. Moreover, the human always plays the first move. So the board will always have an odd number of total marks ('X's and 'O's) where there is one more 'X' than the number of 'O's.

Inputs: List `L` of 'X's, 'O'x, and ' 's

Output: Integer

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
import random

def is_winner(board, letter):
    return  (board[0] == board[1] == board[2] == letter) or \
            (board[3] == board[4] == board[5] == letter) or \
            (board[6] == board[7] == board[8] == letter) or \
            (board[0] == board[3] == board[6] == letter) or \
            (board[1] == board[4] == board[7] == letter) or \
            (board[2] == board[5] == board[8] == letter) or \
            (board[0] == board[4] == board[8] == letter) or \
            (board[2] == board[4] == board[6] == letter)

def possible_moves(board):
    return [key for key, val in enumerate(board) if val == ' ']

def AI_move(L):
    for i in possible_moves(L):
        L[i]='O'
        if is_winner(L,'O')is True:
            return i
        else:
            L[i]=' '
    for i in possible_moves(L):
        L[i]='X'
        if is_winner(L,'X')is True:
            return i
        else:
            L[i]=' '
    result=[]
    for ch in [0,2,6,8]:
        if ch in possible_moves(L):
            result.append(ch)
    if result!=[]:
        return random.choice(result)
    if 4 in possible_moves(L):
        return 4
    return random.choice(possible_moves(L))
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 1.5, grade: true, solution: false -->
```python
assert AI_move(['X',' ','O',' ','X',' ','X',' ','O']) == 5
assert AI_move(['X','X','O',' ','O','X',' ',' ',' ']) == 6
assert AI_move(['X',' ','X',' ','O','X',' ',' ','O']) == 1
assert AI_move(['X','O',' ',' ','O',' ',' ','X','X']) == 6
assert AI_move([' ',' ',' ',' ','X',' ',' ',' ',' ']) in [0,2,6,8]
assert AI_move([' ',' ',' ',' ','X',' ',' ',' ',' ']) in [0,2,6,8]
assert AI_move([' ',' ',' ',' ','X',' ',' ',' ',' ']) in [0,2,6,8]
assert AI_move([' ',' ',' ',' ','X',' ',' ',' ',' ']) in [0,2,6,8]
assert AI_move(['X','O','X',' ',' ','O',' ','X',' ']) in [6,8]
assert AI_move(['X','O','X',' ',' ','O',' ','X',' ']) in [6,8]
```