Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = ""
```

---

# Homework 2
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **Mar 14<sup>th</sup> 2019, 5 pm**

# Problem 1

Given one input variable `n` return the absolute value of `n`.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def absolute_value(n):
    # YOUR CODE HERE
    if n>=0:
        return(n)
    else:
        return(-n)
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert absolute_value(0) == 0
assert absolute_value(123) == 123
assert absolute_value(-98) == 98
```

# Problem 2

Given three input variables `a`, `b`, and `c` return the sum of the largest integer and the smallest integer.

<!-- grade_id: cell-2eb16d0c39243c56, points: 0, grade: false, solution: true -->
```python
def sm_plus_bg(a, b, c):
    # YOUR CODE HERE
    if a<=b<=c or c<=b<=a:
        return(a+c)
    elif a<=c<=b or b<=c<=a:
        return(a+b)
    elif b<=a<=c or c<=a<=b:
        return(b+c)
```

<!-- grade_id: cell-b803298e9a1d8f43, points: 0.5, grade: true, solution: false -->
```python
assert sm_plus_bg(1, 2, 3) == 4
assert sm_plus_bg(3, 2, 1) == 4
assert sm_plus_bg(10, -10, 100) == 90
assert sm_plus_bg(0, 100, 1) == 100
```

# Problem 3

Given a word, return `True` if the word can be typed by the characters which are in same row on a keyboard or `False` otherwise.

<img src="keyboard.png" width="500" align="left">

<!-- grade_id: cell-c7aac4c6c159d799, points: 0, grade: false, solution: true -->
```python
def same_row(W):
    # YOUR CODE HERE
    fr='qwertyuiop'
    sr='asdfghjkl'
    tr='zxcvbnm'
    counter=0
    
    for ch in W:
        if ch.lower()in fr:
            counter+=1
    if counter==len(W):
        return(True)
    else:
        counter=0
        for ch in W:
            if ch.lower()in sr:
                counter+=1
        if counter==len(W):
            return(True)
        else:
            counter=0
            for ch in W:
                if ch.lower()in tr:
                    counter+=1
            if counter==len(W):
                return(True)
            else:
                return (False)
```

```python
def same_row(W):
    kb=['qwertyuiop','asdfghjkl','zxcvbnm']
    for row in kb:
        if W[0]in row:
            h=row
            break
    for i in W:
        if i not in h:
            return False
    return True
        
```

```python
def same_row(W):
    kb=['qwertyuiop','asdfghjkl','zxcvbnm']
    for row in kb:
        if W[0]in row:
            break
    for i in W:
        if i not in row:  # row可以看到其最新的值
            return False  # return就是函数的出口
    return True
```

```python
def same_row(W):
    # YOUR CODE HERE
    fr='qwertyuiop'
    sr='asdfghjkl'
    tr='zxcvbnm'
    counter=0
    if W[0] in fr:
        for i in range(len(W)):
            if W[i] in fr:
                counter+=1
        if counter==len(W):
            return True
        else:
            return False
    if W[0] in sr:
        for i in range(len(W)):
            if W[i] in sr:
                counter+=1
        if counter==len(W):
            return True
        else:
            return False
    if W[0] in tr:
        for i in range(len(W)):
            if W[i] in tr:
                counter+=1
        if counter==len(W):
            return True
        else:
            return False
```

<!-- grade_id: cell-6c59e8bb1e3c37e2, points: 0, grade: true, solution: false -->
```python
assert same_row('hello') == False
assert same_row('alaska') == True
assert same_row('dad') == True
assert same_row('peace') == False
```

```python

```

```python

```