Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = "Yueran Pan"
```

---

# Homework 8
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **March 26<sup>th</sup> 2019, 5 pm**

# Problem 1

The TV series Game of Thrones is very popupar around the world. It has a 9.5 rating on [IMDB](https://www.imdb.com/title/tt0944947/?ref_=fn_al_tt_1). 

The HTML source code of this webpage is saved in GOT-cast.html in the current directory.

Write a function called `get_cast_dict` that extracts the relevant information from the cast list of the HTML file and then creates a dictionary in which the keys are the names of the actresses and actors, and the values are the names of the characters they play.

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
# first version
import re
 
def get_cast_dict():
    d={}
    file=open('GOT-cast.html','r')
    actor=[]
    character=[]
    for line in file:
        if re.search('> ',line):
            actor_name=re.findall('> (\w+ \w+)',line)
            actor.append(actor_name)
        if re.search('&nbsp',line):
            character_name=re.findall('>([ \w\']*)</a>',line)
            character.append(character_name)
    actor_new=[]
    character_new=[]
    for item in actor:
         if item !=[]:
            actor_new=actor_new+item
    for item in character:
         if item !=[]:
            character_new=character_new+item
    del(character_new[-1])

    for i in range(len(actor_new)):
        d[actor_new[i]]=character_new[i]
    return d
```

```python
# after simplification
import re
 
def get_cast_dict():
    d={}
    file=open('GOT-cast.html','r')
    actor=[]
    character=[]
    for line in file:
        actor=actor+re.findall('> (\w+ \w+)',line)
        character=character+re.findall('&nbsp.+>([ \w\']+)</a>',line)
    for i in range(len(actor)):
        d[actor[i]]=character[i]
    return d
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 1, grade: true, solution: false -->
```python
assert get_cast_dict()['Lena Headey'] == 'Cersei Lannister'
assert get_cast_dict()['Sophie Turner'] == 'Sansa Stark'
assert get_cast_dict()['Jack Gleeson'] == 'Joffrey Baratheon'
assert get_cast_dict()['Ben Crompton'] == 'Eddison Tollett'
```

# Problem 2: 

Harry Potter has to learn many magical spells, but it is difficult for him to remember all of the magical spells. He needs your help. 

Write a function called `get_explanation` that uses Harry Potter's notes for magical spells, and returns the explanation when given a magical spell. If you cannot find the magical spell, please return 'what?'

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
# The KEY is the explanation and the VALUE is the magical spell.
magic_notes = {
'the disarming charm': 'expelliarmus',
'send a jet of silver light to hit the enemy': 'rictusempra',
'control the movement of one`s legs': 'tarantallegra',
'shoot a snake out of the end of one`s wand': 'serpensortia',
'light the wand': 'lumos',
'the memory charm': 'obliviate',
'send a Patronus to the dementors': 'expecto patronum',
'the summoning charm': 'accio',
}

def get_explanation(magic_spell):
    for k,v in magic_notes.items():
        if magic_spell==v:
            return k
    return 'what?'
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 1, grade: true, solution: false -->
```python
assert get_explanation('tarantallegra') == 'control the movement of one`s legs'
assert get_explanation('accio') == 'the summoning charm'
assert get_explanation('obliviate') == 'the memory charm'
assert get_explanation('avada kedavra') == 'what?'  # This is The Killing Curse, which is forbidden and unforgivable in the wizarding world.
```