Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = ""
```

---

# Homework 7
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **March 22<sup>nd</sup> 2019, 5 pm**

# Problem 1

Write a function called `create_f_alternator` that takes two functions `f1` and `f2` as input and returns a function that takes an input `n` and returns a string of the results of applying `f1` and `f2` to the integers from 1 to `n` in alternate order.

<pre>
alternator = create_f_alternator(lambda x: x * 2, lambda x: x + 2)
result = alternator(6)
print(result)

2
4
6
6
10
8
</pre>

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
def create_f_alternator(f1, f2):
    def alternator(n):
        L=[]
        for i in range(1,n+1):
            if i%2!=0:
                L.append(str(f1(i)))
            if i%2==0:
                L.append(str(f2(i)))
        result='\n'.join(L)
        return result
    return alternator
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
alternator = create_f_alternator(lambda x: x * 2, lambda x: x + 2)
assert alternator(6) == '2\n4\n6\n6\n10\n8'
alternator = create_f_alternator(lambda x: x ** 2, lambda x: x / 12345)
assert alternator(10) == '1\n0.00016200891049007697\n9\n0.00032401782098015393\n25\n0.0004860267314702309\n49\n0.0006480356419603079\n81\n0.0008100445524503848'
```

# Problem 2

Write a function called `create_digit_remover` which takes an integer `i` between 0 and 9. It returns a function which takes an integer `n` and removes all digits equal to `i`.

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def create_digit_remover(i):
    def remover_i(n):
        result=str(n)
        for ch in result:
            if ch==str(i):
                result=result.replace(ch,'')
        if len(result)==0:
            return 0
        else:
            return int(result)
    return remover_i
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 0.5, grade: true, solution: false -->
```python
remover_2 = create_digit_remover(2)
assert remover_2(2320218) == 3018
assert remover_2(2222) == 0
assert remover_2(134567890) == 134567890

remover_9 = create_digit_remover(9)
assert remover_9(1919191) == 1111
assert remover_9(908742992) == 87422
```

# Problem 3

<img src="derivative.png" width="250" align="left">

Write a function called `create_derivative_function` which takes a function `f` and returns a function that represents the derivative of `f`. The derivative function takes a number `x` as input and returns the rate of change of `f` at `x`.

**hint**: the rate of change of `f` at `x` can be calculated by using a **very** small interval [$x1$, $x2$] where `x` is in the middle of $x1$ and $x2$ and the interval is less than or equal to $0.0002$.

<!-- grade_id: cell-ddc43febbf51524e, points: 0, grade: false, solution: true -->
```python
def create_derivative_function(f):
    def d(x):
        return (f(x+0.0001)-f(x-0.0001))/0.0002 
    return d
```

<!-- grade_id: cell-80aaed74b1fb0813, points: 1, grade: true, solution: false -->
```python
d = create_derivative_function(lambda x: x**2)
assert abs(d(1)-2) < 0.0001
assert abs(d(2)-4) < 0.0001
assert abs(d(10.25)-20.5) < 0.0001

d = create_derivative_function(lambda x: x**4 + x**2 + 1)
assert abs(d(1)-6) < 0.0001
assert abs(d(2.5)-67.5) < 0.0001
assert abs(d(10)-4020) < 0.0001
```

```python

```