Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = ""
```

---

# Homework 1
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **Mar 13<sup>th</sup> 2019, 5 pm**

# Problem 1

Given two input variables `a` and `b` that represent the two sides of right triangle, return the hypotenuse `c` using the Pythagorean Theorem.

$a^2 + b^2 = c^2$

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def hypotenuse(a, b):
    # YOUR CODE HERE
    c=(a**2+b**2)**(1/2)
    return(c)
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert hypotenuse(3, 4) == 5
assert hypotenuse(10,20) == 22.360679774997898
assert hypotenuse(1,100) == 100.00499987500625
```

# Problem 2

Given one input variable `fullname` that is a string, return the family name followed by a comma and the first letter of the given name followed by a period.

<!-- grade_id: cell-150d02900c3a3f16, points: 0, grade: false, solution: true -->
```python
def abbreviation(fullname):
    # YOUR CODE HERE
    A=fullname.split()
    GN=A[0]
    flGN=GN[0]
    FN=A[-1]
    return(FN+','+' '+flGN+'.')
```

<!-- grade_id: cell-58720deb797a5a5e, points: 0.5, grade: true, solution: false -->
```python
assert abbreviation('Daniel F. Lim') == 'Lim, D.'
assert abbreviation('Ming Li') == 'Li, M.'
assert abbreviation('Donald John Trump') == 'Trump, D.'
assert abbreviation('Jinping Xi') == 'Xi, J.'
```

# Problem 3

Given one input variable `S` that contains multiple positive integers separated by one or two colons (at the begining or the end there might be zero or multiple spaces), remove those spaces and colons and return the resulting string as an `int`.

<!-- grade_id: cell-21240d427bd64c8f, points: 0, grade: false, solution: true -->
```python
def get_numbers(S):
    # YOUR CODE HERE
    A=S.replace(':','')
    B=A.replace('::','')
    C=B.strip()
    D=int(C)
    return(D)
```

<!-- grade_id: cell-ed42ea2efa99c3b3, points: 1, grade: true, solution: false -->
```python
assert get_numbers('1:2:3:4:5') == 12345
assert get_numbers('    1::2:3::4:5    ') == 12345
assert get_numbers('12345') == 12345
assert get_numbers(' :  :123::45:  : :') == 12345
```