Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = "Ran Ju"
```

---

# Homework 3
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **March 15<sup>th</sup> 2019, 5 pm**

# Problem 1

A reporter had interviews with many students at DKU. He made notes of these students' personal information. However, his notes are a mess. Could you please standardize his notes?

```python
student1 = 'student1@dukekunshan.edu.cn,23,cHINESE'
student2 = 'student2@dku,datA science,16'
student3 = '17,computer science,student3@DKU'
student4 = 'student4@dukekunshan.edu.cn,29,cHINeSE'
student5 = 'student5@dku,english,16'
student6 = '19,computer science,student6@DKU'
student7 = 'student7@dukekunshan.edu.cn,28,computer SCIENCE'
student8 = 'student8@dku,data science,27'
student9 = '17,COMPUter science,student9@DKU'
```

## Step 1 Get Age

Write a function called `get_age` that takes a string (which represents a student's information) as input and returns the age (as a string).

**hint**: all the ages are 2 digits

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def get_age(student):
    # YOUR CODE HERE
    A=student.split(',')
    B='0123456789'
    for part in A:
        if part[0]in B and part[1] in B:
            return part[0:2]
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert get_age(student1) == '23'
assert get_age(student2) == '16'
assert get_age(student3) == '17'
```

## Step 2 Get E-mail Address

Write a function called `get_email` that returns a student's email address.

**Note** that some email addresses are incorrect, you should return emails with the following suffix: ***@dukekunshan.edu.cn

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def get_email(student):
    # YOUR CODE HERE
    C=student.split(',')
    for part in C:
        if '@'in part:
            name=part.split('@')
            return name[0]+'@dukekunshan.edu.cn'        
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 0.5, grade: true, solution: false -->
```python
assert get_email(student3) == 'student3@dukekunshan.edu.cn'
assert get_email(student4) == 'student4@dukekunshan.edu.cn'
assert get_email(student9) == 'student9@dukekunshan.edu.cn'
```

## Step 3 Get Course

Write a function called `get_course` that returns a student's course. The course should be given in all upper-case letters.

<!-- grade_id: cell-f31c8d516d0460e3, points: 0, grade: false, solution: true -->
```python
def get_course(student):
    # YOUR CODE HERE
    D=student.split(',')
    E='0123456789'
    for part in D:
        if part[0] not in E and '@'not in part:
            return part.upper()
```

<!-- grade_id: cell-ea32695a9bec9a5c, points: 0.5, grade: true, solution: false -->
```python
assert get_course(student7) == 'COMPUTER SCIENCE'
assert get_course(student4) == 'CHINESE'
assert get_course(student8) == 'DATA SCIENCE'
```

## Step 4 Standardize Student

Write a function called `standardize` that applies all functions above and standardizes all the information. All adults' (students 18 years or older) ages should be replaced with '**' to maintain confidentiality. And simply return `None` for those who study data science.

The result should be like:

student1: student1@dukekunshan.edu.cn,**,CHINESE

student2: None

student3: student3@dukekunshan.edu.cn,17,COMPUTER SCIENCE

<!-- grade_id: cell-804725bd9bc5f2ad, points: 0, grade: false, solution: true -->
```python
def standardize(student):
    # YOUR CODE HERE
    M=student.upper()
    N='DATA SCIENCE'
    if N in M:
        return None
    else:
        if int(get_age(student))>=18:
            return get_email(student)+','+'**'+','+get_course(student).upper()
        else:
            return get_email(student)+','+get_age(student)+','+get_course(student).upper()
```

<!-- grade_id: cell-1f0b7468d08c02e9, points: 0.5, grade: true, solution: false -->
```python
assert standardize(student1) == 'student1@dukekunshan.edu.cn,**,CHINESE'
assert standardize(student2) == None
assert standardize(student3) == 'student3@dukekunshan.edu.cn,17,COMPUTER SCIENCE'
```

```python

```