Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = ""
```

---

# Homework 5
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **March 20<sup>th</sup> 2019, 5 pm**

# Problem 1

Given two lists `L1` and `L2`, where each list consists of 2-dimensional coordinates $(x_1, x_2)$, return a merged list sorted in increasing order according to the Euclidean distance between the 2-dimensional coordinate $(x_1, x_2)$ and the origin $(0, 0)$.

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
def mergeList(L1, L2):
    list=[]
    for tuple in L1:
        list.append(tuple)
    for tuple in L2:
        list.append(tuple)
    def distance(tuple):
        return (tuple[0])**2+(tuple[1])**2
    result=sorted(list, key=distance) 
    return result
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 1, grade: true, solution: false -->
```python
assert mergeList([(3, 3), (2, 2)], [(6, 6), (0, 0)]) == [(0, 0), (2, 2), (3, 3), (6, 6)]
assert mergeList([(5, 1), (2, 3)], [(7, 6), (1, 4)]) == [(2, 3), (1, 4), (5, 1), (7, 6)]
assert mergeList([(-3.4, 8.9), (0.8, 3.4)], [(5.5, -2.3), (-6.4, -9.7)]) == [(0.8, 3.4), (5.5, -2.3), (-3.4, 8.9), (-6.4, -9.7)]
```

# Problem 2: 

### Preparation

The file CA-GrQc.txt from the Stanford Large Network Dataset Collection is about co-authorship links for articles in the ArXiv category General Relativity. Each number represents an author.

### Get each author's coauthors  
Enter each author's unique coauthors in a dictionary and return that dictionary. The dictionary should look something like: `{13: [7596, 11196, 19170], 14:[14171], ...}`.

#### Hints

Notice that the data contains errors. For example, the data shows that author 13 coauthored with himself/herself, which is meaningless. Your code should exclude oneself in one's list of coauthors.

#### Note   

You need to convert the number from `str` type to `int` type, since reading from a file always works with `str`.

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def coauthors_dict():
    coauthors = dict()

    file=open('CA-GrQc.txt','r')
    cooperation=[]
    authors=[]
    for line in file:
        if line[0]!='#':
            num=line.split()
            L=[int(num[0]),int(num[1])]
            cooperation.append(L)
            authors.append(int(num[0]))
            authors.append(int(num[1]))               
    unique_authors=list(set(authors))
#     for name in authors:
#         if name not in unique_authors:
#             unique_authors.append(name)
    for name in unique_authors:
        colist=[]
        for couple in cooperation:
            if name==couple[0]:
                colist.append(couple[1])
            #if name==couple[1]:
                #colist.append(couple[0])
        
        coauthors[name]=sorted(colist)
        #print(colist)
  
    return coauthors
```

```python
def coauthors_dict():
    coauthors = {}
    cooperation=[]
    file=open('CA-GrQc.txt','r')
    for line in file:
        if line[0]!='#':
            num=line.split()
            A=int(num[0])
            B=int(num[1])
            L=[A,B]
            cooperation.append(L)
    for A,B in cooperation:
        if A not in coauthors:
            coauthors[A]=[B]
        else:
            coauthors[A].append(B)
    return coauthors        
```

```python
def coauthors_dict():
    a=open('CA-GrQc.txt')
    coauthor={}
    d=[]
    c=[]
    for line in a:
        if '#' in line:
            line=''
        lines=line.split()
        b=[]
        for i in lines:
            i=int(i)
            b.append(i)
        c.append(b)
    c=c[4:]
    for b in c:
        if b[0] not in d:
            d.append(b[0])
    d=sorted(d)
    for i in d:
        if i not in coauthor:
            coauthor[i]=[]
    for b in c:
        if b[0] in coauthor:
            coauthor[b[0]].append(b[1])
    return coauthor
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 1, grade: true, solution: false -->
```python
assert coauthors_dict()[14] == [14171]
assert len(coauthors_dict()[45]) == 49
assert coauthors_dict()[26] == [1407, 4550, 11801, 13096, 13142]
assert coauthors_dict()[24] == [3858, 15774, 19517, 23161]
```