Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = ""
```

---

# Homework 6
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **March 21<sup>th</sup> 2019, 5 pm**

# Problem 1

Write a function called `power` that takes two integers `x` and `n` and returns $x^n$ using **recursion**.

**Hint**:

Do not use `for` or `while` loops, since it costs too much time.

if `n` is even, $x^n = x^{\frac{n}{2}} x^{\frac{n}{2}}$

if `n` is odd, $x^n = x^{\frac{n}{2}} x^{\frac{n}{2}} x$

<!-- grade_id: cell-46c55fb5434e1162, points: 0, grade: false, solution: true -->
```python
def power(x, n):
    if n==1:
        return x
    if n%2==0:
        return(power(x, n/2))**2
    if n%2!=0:
        return((power(x, (n-1)/2))**2)*x
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 1, grade: true, solution: false -->
```python
assert power(100, 2) == 10000
assert power(1, 10000000) == 1
assert power(2, 10) == 1024
assert power(11, 11) == 285311670611
```

# Problem 2: 

It is time for Jim to go home but the elevator is out of order. Jim has to climb a staircase with `n` steps. Each time he can either climb one or two steps.

Write a function called `num_ways` that takes an integer `n` as input and returns the number of distinct ways Jim he reach the top of a staircase with `n` steps.

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def num_ways(n):
    if n==1:
        return 1
    if n==2:
        return 2
    return num_ways(n-1)+num_ways(n-2)
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 1, grade: true, solution: false -->
```python
assert num_ways(2) == 2
assert num_ways(5) == 8
assert num_ways(15) == 987
assert num_ways(20) == 10946
```