Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = ""
```

---

# Lab 1
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 14<sup>th</sup> 2019**

# Problem 1

Write a function `count_vowels` that takes a string `S` as input and returns the number of vowels that occur in `S`.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def count_vowels(S):
    A='aeiou'
    counter=0
    for ch in S:
        if ch.lower() in A:
            counter+=1
    return counter
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert count_vowels('Evander Holyfield') == 6
assert count_vowels('AI') == 2
assert count_vowels('bcdfghjklmnpqrstvwxyz') == 0
assert count_vowels('') == 0
```

# Problem 2

Write a function called `divisors` that takes an integer `n` as input and returns a string of all the integers that evenly divide `n` (besides 1 and `n`).

<!-- grade_id: cell-150d02900c3a3f16, points: 0, grade: false, solution: true -->
```python
def divisors(n):
    result=''
    for i in range(2,n):
        if n%i==0:
            result=result+str(i)+' '
    A=result.strip()
    return A
```

<!-- grade_id: cell-58720deb797a5a5e, points: 0.5, grade: true, solution: false -->
```python
assert divisors(10) == '2 5'
assert divisors(36) == '2 3 4 6 9 12 18'
assert divisors(1000) == '2 4 5 8 10 20 25 40 50 100 125 200 250 500'
assert divisors(2001) == '3 23 29 69 87 667'
assert divisors(567890) == '2 5 10 109 218 521 545 1042 1090 2605 5210 56789 113578 283945'
```

# Problem 3

Write a function called `exists` that takes two strings `letters` and `S`. It returns `True` if all the letters in `letters` exist in `S` and `False` otherwise.

<!-- grade_id: cell-21240d427bd64c8f, points: 0, grade: false, solution: true -->
```python
def exists(letters, S):
    counter=0
    for ch in letters:
        if ch.lower() in S.lower():
            counter+=1
    if counter==len(letters):
        return True
    else:
        return False
```

<!-- grade_id: cell-ed42ea2efa99c3b3, points: 1, grade: true, solution: false -->
```python
assert exists('abr','Barack') == True
assert exists('abr','Hello World') == False
assert exists('','Hello World') == True
assert exists('ABC','abc') == True
```

# Problem 4

Write a function called `match` that takes two strings `S1` and `S2` as inputs and returns a string of all the letters that exist in both strings.

**note**: you must have a working version of function `exists` in order to test this function.

<!-- grade_id: cell-77445544dd3fefa5, points: 0, grade: false, solution: true -->
```python
def match(S1, S2):
    A=''
    for ch in S1:
        if ch in S2:
            A=A+ch
    return A
```

<!-- grade_id: cell-8b21a078a98867c2, points: 0, grade: true, solution: false -->
```python
m = match('Hello World','Barack Obama')
assert exists(m,'Hello World')
assert exists(m,'Barack Obama')

m = match('Hello World','Hello')
assert exists(m,'Hello World')
assert exists(m,'Hello')
```

# Problem 5

Write a function called `movie` that has no inputs. It uses the `input` function (with string: 'How old are you? ') to get the **age** of the user and prints information about what movies the user can watch.

If the age is greater than 17, then you should print 'You can watch R, PG-13, and PG movies'.

If the age is greater than 12 but less than or equal to 17, then you should print 'You can watch PG-13 and PG movies, but NOT R movies'.

If the age is less than or equal to 12, then you should print 'You can only watch PG movies'.

<!-- grade_id: cell-7dad4e2e8b4e690a, points: 0, grade: false, solution: true -->
```python
def movie():
    age=input('How old are you?')
    if int(age)>17:
        print('You can watch R, PG-13, and PG movies')
    elif 12<int(age)<=17:
        print('You can watch PG-13 and PG movies, but NOT R movies')
    elif int(age)<=12:
        print('You can only watch PG movies')
```