Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Lab 4
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**April 4<sup>th</sup> 2019**

# Problem 1

A list can be used to implement a `stack`. A `stack` is a data structure that ensures first-in-last-out (FILO) behavior. So if you insert (technically called a `push`) values into a stack you will retrieve (technically called a `pop`) the values in reverse order.

Write two functions `push` and `pop`.

`push` takes two inputs: the list `L` that represents the stack and an integer `val` that is being inserted into `L`. This function then returns `L`.

`pop` takes the list `L` as input and removes the integer that was last inserted from `L` and returns this integer.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def push(L, val):
    L.append(val)

def pop(L):
    last=L[-1]
    del(L[-1])
    return last
```

<!-- grade_id: cell-764b2aa492773005, points: 1, grade: true, solution: false -->
```python
stack = []
push(stack, 1)
push(stack, 2)
push(stack, 3)
assert pop(stack) == 3
push(stack, 4)
push(stack, 5)
assert pop(stack) == 5
assert stack == [1,2,4]
```

# Problem 2

The Stock Span problem is a financial problem where we have a list of `n` daily price quotes for a stock and we need to calculate the spans of the stock's price for all `n` days.

The span `S`<sub>i</sub> of the stock’s price on day `i` is defined as the maximum number of consecutive days before the `i`th day, for which the price of the stock on the earliest day is less than the price on the `i`th day.

For example, if a list of 7 days of prices is given as [100, 80, 60, 70, 60, 75, 85], then the span values for the corresponding 7 days are [1, 1, 1, 2, 1, 4, 6].

<!-- grade_id: cell-77b8dd1a59eda2ee, points: 0, grade: false, solution: true -->
```python
def stock_span(L):
    result=[1]
    if len(L)>=2:
        for i in range(1,len(L)):
            counter=1
            index=i-1
            while L[i]>L[index]:
                counter=counter+1
                index=index-1
                if index==-1:
                    break
            result.append(counter)
    return result
```

<!-- grade_id: cell-209c42fabacb044c, points: 1, grade: true, solution: false -->
```python
import random

def gen_list(n):
    return [random.randint(0,100000) for i in range(n)]

assert stock_span([100]) == [1]
assert stock_span([100,80,60,70,60,75,85]) == [1,1,1,2,1,4,6]
assert stock_span([97, 65, 94, 199, 160, 139, 136, 108, 19, 166, 75, 160, 196, 105, 99, 39, 102, 40, 44, 13, 84, 156, 94, 91, 141, 197, 21, 10, 170, 171, 48, 147, 51, 127, 43, 68, 80, 101, 38, 142, 27, 83, 85, 162, 157, 184, 25, 30, 162, 162]) == [1, 1, 2, 4, 1, 1, 1, 1, 1, 6, 1, 2, 9, 1, 1, 1, 3, 1, 2, 1, 4, 9, 1, 1, 3, 22, 1, 1, 3, 4, 1, 2, 1, 2, 1, 2, 3, 4, 1, 8, 1, 2, 3, 14, 1, 20, 1, 2, 3, 1]
result = %timeit -o stock_span(sorted(gen_list(1000)))
```

Output:
```
70.1 ms ± 482 µs per loop (mean ± std. dev. of 7 runs, 10 loops each)
```

# Problem 3

A more efficient solution exists that utilizes the functionality of a `stack`. The idea is to `push` the index of the first stock price into the `stack`.

Then we can go through the following steps for the rest of the stock prices:

1. For the stock price at index `i`, compare it with the stock price at index `j` where `j` is the value at the top of the `stack`.
2. If the stock price at `i` is less than or equal to the stock price at `j` then `pop` the stack. Continue doing this until the stock price at `i` is greater than `j`.
3. Now we know that the stock price at index `j` is greater than the stock price at index `i` so the `span` for the stock price at `i` will be `i` - `j`.
4. If this is not the case then the stack is empty and there is no other previous stock price that is greater than the stock price at index `i` so the `span` for the stock price at `i` will be `i` + 1.

<!-- grade_id: cell-6beaeece872ef205, points: 0, grade: false, solution: true -->
```python
def stock_span_stack(L):
    def helper(i):
        if len(stack)==0:
            push(stack,i)
            return i+1
        else:
            j=stack[-1]
            while L[i]>L[j]:
                pop(stack)
                if len(stack)==0:
                    push(stack,i)
                    return i+1
                else:
                    j=stack[-1]
                    continue
            push(stack,i)
            return i-j
            
        
    result=[1]
    stack=[0]
    for i in range(1,len(L)):
        result.append(helper(i))
    return result
```

<!-- grade_id: cell-bcc453928b6b56b1, points: 1, grade: true, solution: false -->
```python
assert stock_span_stack([100]) == [1]
assert stock_span_stack([100,80,60,70,60,75,85]) == [1,1,1,2,1,4,6]
assert stock_span_stack([97, 65, 94, 199, 160, 139, 136, 108, 19, 166, 75, 160, 196, 105, 99, 39, 102, 40, 44, 13, 84, 156, 94, 91, 141, 197, 21, 10, 170, 171, 48, 147, 51, 127, 43, 68, 80, 101, 38, 142, 27, 83, 85, 162, 157, 184, 25, 30, 162, 162]) == [1, 1, 2, 4, 1, 1, 1, 1, 1, 6, 1, 2, 9, 1, 1, 1, 3, 1, 2, 1, 4, 9, 1, 1, 3, 22, 1, 1, 3, 4, 1, 2, 1, 2, 1, 2, 3, 4, 1, 8, 1, 2, 3, 14, 1, 20, 1, 2, 3, 1]
result = %timeit -o stock_span_stack(sorted(gen_list(1000)))
```

Output:
```
2.52 ms ± 57.1 µs per loop (mean ± std. dev. of 7 runs, 100 loops each)
```