Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Lab 5
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**April 11<sup>th</sup> 2019**

# Problem 1

Write a recursive function called `display` that takes a node `n` (which is the head of a linked list) as input and returns a string of all the values of the nodes of the linked list starting at `n` in reverse order.

**note**: do not use a `for` loop or a `while` loop!

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
class node(object):
    def __init__(self, val):
        self.val = val
        self.next = None
        
def display(n):
    result=''
    if n.next==None:
        return str(n.val)
    else:
        return display(n.next)+str(n.val)
```

<!-- grade_id: cell-764b2aa492773005, points: 1, grade: true, solution: false -->
```python
n1 = node(1)
n2 = node(2)
n3 = node(3)
n4 = node(4)
n5 = node(5)

n1.next = n2
n2.next = n3
n3.next = n4
n4.next = n5

assert display(n1) == '54321'
assert display(n3) == '543'
assert display(n5) == '5'
```

# Problem 2

Write a function called `count_nodes` that takes a node `n` that represents the root of a binary tree and an integer `level` that represents the level of traversal. This function returns the number of nodes that are at `level`.

<!-- grade_id: cell-77b8dd1a59eda2ee, points: 0, grade: false, solution: true -->
```python
class node(object):
    def __init__(self, val, lt=None, rt=None):
        self.val = val
        self.lt = lt
        self.rt = rt

def count_nodes(n, level):
    pack=[[n]]
    for i in range(level+1):
        sub_pack=[]
        for ch in pack[i]:
            if ch.lt:
                sub_pack.append(ch.lt)
            if ch.rt:
                sub_pack.append(ch.rt)
        pack.append(sub_pack)
    return len(pack[level])
```

<!-- grade_id: cell-209c42fabacb044c, points: 1, grade: true, solution: false -->
```python
n1 = node(1)
n2 = node(2)
n3 = node(3)
n4 = node(4)
n5 = node(5)
n6 = node(6)
n7 = node(7)
n8 = node(8)
n9 = node(9)
n10 = node(10)
n1.lt = n2
n1.rt = n3
n2.lt = n4
n2.rt = n5
n3.lt = n6
n3.rt = n7
n4.lt = n8
n4.rt = n9
n8.lt = n10

assert count_nodes(n1,0) == 1
assert count_nodes(n1,1) == 2
assert count_nodes(n1,2) == 4
assert count_nodes(n1,3) == 2
assert count_nodes(n1,4) == 1
```

# Problem 3

Write a function called `delete_dups` that takes a node `n` (that represents the head of the linked list) and removes all the duplicate values in the linked list.

**Note**: duplicate values only exist in contiguous groups. That is, no duplicate values will exist separately. You can assume that only valid linked lists will be given as inputs to this function.

<pre>
[1, 1, 1, 2, 3, 4, 4, 5]  -> valid
[1, 1, 2, 1, 3, 3, 4, 4]  -> invalid
</pre>

<!-- grade_id: cell-6beaeece872ef205, points: 0, grade: false, solution: true -->
```python
class node(object):
    def __init__(self, val):
        self.val = val
        self.next = None
        
def display(n):
    result = ''
    while n is not None:
        result += str(n.val)
        n = n.next
    return result

def delete_dups(n):
    preN=n
    curN=n.next
    while curN!=None:
        while curN!=None and preN.val==curN.val:
            preN.next=curN.next
            curN=curN.next
        if curN!=None:
            preN=curN
            curN=curN.next
```

<!-- grade_id: cell-bcc453928b6b56b1, points: 1, grade: true, solution: false -->
```python
n1 = node(1)
n2 = node(1)
n3 = node(1)
n4 = node(2)
n5 = node(2)

n1.next = n2
n2.next = n3
n3.next = n4
n4.next = n5

assert display(n1) == '11122'
delete_dups(n1)
assert display(n1) == '12'

n6 = node(3)
n4.next = n6

assert display(n1) == '123'

n7 = node(2)
n8 = node(2)
n9 = node(2)
n10 = node(1)

n6.next = n7
n7.next = n8
n8.next = n9
n9.next = n10

assert display(n1) == '1232221'
delete_dups(n1)
assert display(n1) == '12321'
```

```python

```