Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Lab 2
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 21<sup>th</sup> 2019**

# Problem 1

Write a function `all_distinct` that takes a list of integers `L` as input and returns `True` if all the integers are different from each other and `False` otherwise.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def all_distinct(L):
    List=list(set(L))
    if List==L:
        return True
    else:
        return False
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert all_distinct([1,2,3,4,5]) == True
assert all_distinct([]) == True
assert all_distinct([1]) == True
assert all_distinct([1,1,1,1,1]) == False
assert all_distinct([1,2,3,4,1]) == False
```

# Problem 2

Write a function called `remove_n_create` that takes a list of integers `L` and an integer `n`. Beginning with the 1<sup>st</sup> integer this function will create a new integer with every `n`<sup>th</sup> integer as its next digit.

<pre>
remove_n_create([1,2,3,4,5,6],2) -> 135264
</pre>

<!-- grade_id: cell-150d02900c3a3f16, points: 0, grade: false, solution: true -->
```python
def remove_n_create(L, n):
    result=[str(L[0])]
    del(L[0])
    while L!=[]:
        try:
            result.append(str(L[n-1]))
            del(L[n-1])
        except:
            result.append(str(L[n-1]))
    return int(''.join(result))        
```

<!-- grade_id: cell-58720deb797a5a5e, points: 0.5, grade: true, solution: false -->
```python
assert remove_n_create([1,2,3,4,5,6],1) == 123456
assert remove_n_create([1,2,3,4,5,6],2) == 135264
assert remove_n_create([1,2,3,4,5,6],3) == 142635
assert remove_n_create([1,2,3,4,5,6],4) == 154632
assert remove_n_create([1,2,3,4,5,6],5) == 162453
assert remove_n_create([1,2,3,4,5,6],6) == 124365
```

Output:
```
<ipython-input-16-32bc9dc57c49> in <module>()
      1 assert remove_n_create([1,2,3,4,5,6],1) == 123456
----> 2 assert remove_n_create([1,2,3,4,5,6],2) == 135264
      3 assert remove_n_create([1,2,3,4,5,6],3) == 142635
      4 assert remove_n_create([1,2,3,4,5,6],4) == 154632
      5 assert remove_n_create([1,2,3,4,5,6],5) == 162453

<ipython-input-15-cec7a566c708> in remove_n_create(L, n)
      3     del(L[0])
      4     while L!=[]:
----> 5         result.append(str(L[n-1]))
      6         del(L[n-1])
      7     return int(''.join(result))

IndexError: list index out of range
```

# Problem 3

Write a function called `strobo` that takes an integer `n` and returns a list of all strobogrammatic numbers as **strings** that are of length `n`. 

**note**: The order of these strings does not matter.

A strobogrammatic number is a number whose numeral is rotationally symmetric, so that it appears the same when rotated 180 degrees. In other words, the numeral looks the same right-side up and upside down (e.g., 69, 96, 1001).

<img src="strobo.png" width="500">

<p>

<pre>
strobo(1) -> ['1', '0', '8']
strobo(2) -> ['11', '69', '88', '96']
strobo(3) -> ['818', '111', '916', '619', '808', '101', '906', '609', '888', '181', '986', '689']
</pre>

<!-- grade_id: cell-21240d427bd64c8f, points: 0, grade: false, solution: true -->
```python
def strobo(n):
    sym_digits=['0','1','8']
    sym_pairs=['9','6']
    
```

<!-- grade_id: cell-ed42ea2efa99c3b3, points: 1, grade: true, solution: false -->
```python
assert sorted(strobo(1)) == ['0', '1', '8']
assert sorted(strobo(2)) == ['11', '69', '88', '96']
assert sorted(strobo(3)) == ['101', '111', '181', '609', '619', '689', '808', '818', '888', '906', '916', '986']
assert sorted(strobo(4)) == ['1001', '1111', '1691', '1881', '1961', '6009', '6119', '6699', '6889', '6969', '8008', '8118', '8698', '8888', '8968', '9006', '9116', '9696', '9886', '9966']
```

# Problem 4

Write a function called `subtract_itself` that takes an integer `n` and subtracts the sum of the digits of `n` from `n`.

<!-- grade_id: cell-77445544dd3fefa5, points: 0, grade: false, solution: true -->
```python
def subtract_itself(n):
    for i in str(n):
        n=n-int(i)
    return n
```

<!-- grade_id: cell-8b21a078a98867c2, points: 0, grade: true, solution: false -->
```python
assert subtract_itself(4) == 0
assert subtract_itself(10) == 9
assert subtract_itself(56) == 45
assert subtract_itself(1291) == 1278
```

# Problem 5

Write a function called `perfect_squares` that takes a list of integers as input and returns a list of only those integers that are perfect squares.

<!-- grade_id: cell-7dad4e2e8b4e690a, points: 0, grade: false, solution: true -->
```python
from math import sqrt

def perfect_squares(L):
    result=[]
    for num in L:
        if (sqrt(num)-int(sqrt(num)))==0:
            result.append(num)
    return result
```

```python
def perfect_squares(L):
    square=[i**2 for i in range(100)]
    result=[]
    for num in L:
        if num in square:
            result.append(num)
    return result
```

<!-- grade_id: cell-b526b7d3397a5397, points: 0, grade: true, solution: false -->
```python
assert perfect_squares([1,2,3,4,5,9,12,13,14,16,22,25]) == [1,4,9,16,25]
```