Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# MiniProject 2
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **Mar 25<sup>th</sup> 2019, 5 pm**

# Problem 1

Assume that a valid phone number must appear in one of the following two formats: (xxx) xxx-xxxx or xxx-xxx-xxxx. (x is a digit).

Write a function called `valid_number` that takes a string `phone_number` and returns `True` if it is a valid phone number and `False` otherwise.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def valid_number(numstr):
    num='1234567890'
    
    if len(numstr)==12:
        for ch in (numstr[:3]and numstr[4:7]and numstr[8:12]):
            if ch not in num:
                return False
        if numstr[3] !='-' or numstr[7] !='-':
                return False
        else:
            return True
    if len(numstr)==14:
        for ch in (numstr[1:4]and numstr[6:9]and numstr[10:14]):
            if ch not in num:
                return False
        if numstr[0]!='(' or numstr[4]!=')':
            return False
        if numstr[5]!=' ':
            return False
        if numstr[9]!='-':
            return False
        else:
            return True
    else:
        return False
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 1, grade: true, solution: false -->
```python
assert valid_number(' 287-323-4567') == False
assert valid_number('321-456-1213') == True
assert valid_number('213-1234-1213') == False
assert valid_number('(871)716-1243') == False
assert valid_number('(672) 716-1234') == True
assert valid_number('817-2139976') == False
assert valid_number('1123-121-8876') == False
assert valid_number(' (672) 716-1234') == False
assert valid_number('(672) 716-1234 ') == False
```

# Problem 2

The `count and say` sequence is the sequence of integers where the first six terms are as follows:

<ol>
<li>1
<li>11
<li>21
<li>1211
<li>111221
<li>312211
</ol>

`1` is read off as `"one 1"` or `11`.  
`11` is read off as `"two 1s"` or `21`.  
`21` is read off as `"one 2, then one 1"` or `1211`

Write a function called `count_and_say` that takes an integer `n` (1 $\leq$ n $\leq$ 30) as input and returns the $n^{th}$ term of the count and say sequence.

__Note__: Each term of the sequence of integers will be represented as a string.

<!-- grade_id: cell-150d02900c3a3f16, points: 0, grade: false, solution: true -->
```python
def count_and_say(n):
    if n==1:
        return '1'
    else:
        L=count_and_say(n-1)
        A=[]
        for i in range(0,len(L)-1):
            A.append(L[i])
            if L[i]!=L[i+1]:
                A.append(' ')
        A.append(L[-1])       
        B=''.join(A)
        C=B.split()
        result=''
        for item in C:
            result=result+str(len(item))+item[0]
        return result
```

<!-- grade_id: cell-58720deb797a5a5e, points: 1, grade: true, solution: false -->
```python
assert count_and_say(3) == '21'
assert count_and_say(6) == '312211'
assert count_and_say(12)[-6:] == '113211'
```

# Problem 3

Write a function called `find_pairs` that takes a list of integers `L` and an integer `k` and returns the number of **unique** `k`-diff pairs. A `k`-diff pair is defined as an integer pair (i, j), where `i` and `j` are both numbers in the list and their absolute difference is `k`.

#### Example 1
**Input:** [3, 1, 4, 1, 5], k = 2   
**Output:** 2  
**Note:** There are two 2-diff pairs in the array, (1, 3) and (3, 5).
Although we have two 1s in the input, we should only return the number of **unique** pairs.

#### Example 2
**Input:** [1, 2, 3, 4, 5], k = 1  
**Output:** 4  
**Note:** There are four 1-diff pairs in the array, (1, 2), (2, 3), (3, 4) and (4, 5).

#### Example 3
**Input:** [1, 3, 1, 5, 4], k = 0  
**Output:** 1  
**Note:** There is one 0-diff pair in the array, (1, 1).

<!-- grade_id: cell-21240d427bd64c8f, points: 0, grade: false, solution: true -->
```python
def find_pairs(L, k):
    List=list(set(L))
    counter=0
    for num1 in List:
        for num2 in List:
            if num1-num2==k:
                counter=counter+1
    if k!=0:
        return counter
    if k==0:
        if L==List:
            return counter-len(List)
        else:
            return counter-len(List)+1
```

<!-- grade_id: cell-ed42ea2efa99c3b3, points: 1, grade: true, solution: false -->
```python
assert find_pairs([1, 2, 3, 4, 5],1) == 4
assert find_pairs([3, 1, 4, 1, 5],2) == 2
assert find_pairs([1, 3, 1, 5, 4],0) == 1
assert find_pairs([3, 1, 3, 6, 2, 8, 9, 11, 4, 1, 5], 4) == 4
assert find_pairs([3, 1, 3, 6, 6, 3, 9, 3, 9, 3, 1], 3) == 2
assert find_pairs([1, 1], 1) == 0
assert find_pairs([1, 3, 7, 5, 11, 9, 8, 6, 3, 4, 9, 1, 2, 1], 2) == 8
assert find_pairs([7, 5, 9, 8, 6, 3, 4, 9, 11, 2, 1, 9], 2) == 8
```

```python

```