Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# MiniProject 4
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 8<sup>th</sup> 2019, 5 pm**

# Problem 1

A set of parentheses is balanced if each left parenthesis has an appropriately matching right parenthesis. The following sets of parentheses are balanced:

<pre>
(((())))
(()(())()()())
</pre>

But the following sets of parentheses are not balanced:

<pre>
(()()))(
(()))())
</pre>

Write a function called `is_balanced` that takes a string `S` of parentheses and returns `True` if `S` is balanced and `False` otherwise.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def is_balanced(S):
    stack=[]
    for ch in S:
        if ch=='(':
            stack.append(ch)
        if ch==')':
            try:
                del(stack[-1])
            except:
                return False
    if stack!=[]:
        return False
    return True
```

<!-- grade_id: cell-15b2323f12f4dfb0, points: 1, grade: true, solution: false -->
```python
assert is_balanced('()') == True
assert is_balanced('())') == False
assert is_balanced('(()') == False
assert is_balanced('(((())))') == True
assert is_balanced('(()()))(') == False
assert is_balanced('(()))())') == False
assert is_balanced('(()(())()()())') == True
assert is_balanced('((()()(()))())') == True
assert is_balanced('((()((())())))') == True
assert is_balanced('((())(()(()())))()()((()()))') == True
assert is_balanced('((())(()(()())))()()((())()))') == False
```

# Problem 2

Write a function called `calculator` that takes a string `E` that represents an arithmetic expression (involving parentheses, integers, variables, and arithemtic operators (+, -, \*, /)) and a dictionary `dict` that represents a mapping of variables to integers. This will perform the relevant operations while preserving the precedence of parentheses over operators.

**note**: you can assume that the arithmetic expression has balanced parentheses.

<!-- grade_id: cell-293f28079a5e80a7, points: 0, grade: false, solution: true -->
```python
def calculator(E, dict):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d4032e88631e55bc, points: 2, grade: true, solution: false -->
```python
assert calculator('8 * (   5 -    1)',{}) == 32
assert calculator('8*(5-1)',{}) == 32
assert calculator('(8*(5-1)) + 1',{}) == 33
assert calculator('((a - b) * c)',{'a':2,'b':3,'c':4}) == -4
assert calculator('((a - 4) / (b + 3))',{'a':12,'b':31,'c':14}) == 0.23529411764705882
assert calculator('((4 - a) / (b + 3)) + c',{'a':12,'b':31,'c':14}) == 13.764705882352942
```