Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Xuchen Gong"
COLLABORATORS = "Ran Ju"
```

---

# MiniProject 1
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **Mar 18<sup>th</sup> 2019, 5 pm**

# Problem 1

Write a function called `reverse_words` that takes string `S` as input and returns a string with each of the substrings with only characters in the alphabet reversed.

__hint__:

S.isalpha() will return true if all the characters in an input string are letters in the alphabet.

<pre>
s1 = 'a'
s1.isalpha() -> True

s2 = '-'
s2.isalpha() -> False
</pre>

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def reverse_words(s):
    unalpha=[]
    for i in range(len(s)):
        if s[i].isalpha()==False:
            unalpha.append(i)
    unalpha.append(len(s))
    result=''
    a=s[:unalpha[0]][::-1]     # the first reversed part
    for n in range(0,len(unalpha)-1):
        result+=s[unalpha[n]]+s[unalpha[n]+1:unalpha[n+1]][::-1]
    return ( a+result)
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert reverse_words('duke-kunshan-university') == 'ekud-nahsnuk-ytisrevinu'
assert reverse_words('computer science') == 'retupmoc ecneics'
assert reverse_words('***@dukekunshan.edu.cn') == '***@nahsnukekud.ude.nc'
assert reverse_words('Q!W@E#R$T%Y^U&I*O(P)') == 'Q!W@E#R$T%Y^U&I*O(P)'
```

# Problem 2

A ball falls from `H` meters above the ground and bounces back to half the previous height. Then it drops again and bounces back again to half the previous height. Write a function called `fall` that returns how far the ball has traveled when it falls `N` times (`N` > 0) from height `H` and also return the height it will bounce back to?

<!-- grade_id: cell-150d02900c3a3f16, points: 0, grade: false, solution: true -->
```python
def fall(H, N):
    has_traveled=float(H)*(3-2**(2-float(N)))
    will_bounce_to=float(H)*2**(-float(N))
    return has_traveled,will_bounce_to 
```

<!-- grade_id: cell-58720deb797a5a5e, points: 0.5, grade: true, solution: false -->
```python
assert fall(50, 2) == (100.0000000000, 12.5000000000)
assert fall(50, 4) == (137.5000000000, 3.1250000000)
assert fall(50, 6) == (146.8750000000, 0.7812500000)
assert fall(50, 8) == (149.2187500000, 0.1953125000)
assert fall(100, 2) == (200.0000000000, 25.0000000000)
assert fall(100, 4) == (275.0000000000, 6.2500000000)
assert fall(100, 6) == (293.7500000000, 1.5625000000)
assert fall(100, 8) == (298.4375000000, 0.3906250000)
assert fall(150, 2) == (300.0000000000, 37.5000000000)
assert fall(150, 4) == (412.5000000000, 9.3750000000)
assert fall(150, 6) == (440.6250000000, 2.3437500000)
assert fall(150, 8) == (447.6562500000, 0.5859375000)
assert fall(10, 2) == (20.0000000000, 2.5000000000)
assert fall(10, 4) == (27.5000000000, 0.6250000000)
assert fall(10, 6) == (29.3750000000, 0.1562500000)
assert fall(10, 8) == (29.8437500000, 0.0390625000)
```

# Problem 3

A palindromic number or numeral palindrome is a number that remains the same when its digits are reversed. It is symmetrical ike 16461.

Other examples include: 121, 3, 76722762, 33344333

Write a function called `find_pal` that takes an integer `n` as input and returns the smallest palindromic number that is bigger than `n`.

<!-- grade_id: cell-21240d427bd64c8f, points: 0, grade: false, solution: true -->
```python
def find_pal(num):
    for i in range(num+1,100000000):
        if str(i)[:]==str(i)[::-1]:
            return i
```

<!-- grade_id: cell-ed42ea2efa99c3b3, points: 1, grade: true, solution: false -->
```python
assert find_pal(3) == 4
assert find_pal(131) == 141
assert find_pal(2447) == 2552
assert find_pal(12141) == 12221
```