Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# MiniProject 5
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

Due: **April 15<sup>th</sup> 2019, 5 pm**

A **Balanced Binary Search Tree** is a binary tree that is constrained by the following two properties:

1. The value of a node will always be greater than the value of its left child and less than the value of its right child.

2. A given node's left child and right child will have heights with an absolute difference of 0 or 1.

These properties will guarantee that searching for an existing value in the search tree will take logarithmic time.

```python
# Binary Search Tree Node
class bst_node(object):
    def __init__(self, val, lt=None, rt=None, pt=None):
        self.val = val
        self.lt = lt
        self.rt = rt
        self.pt = pt
        self.height = 1

# returns a string of the values in the BST
# according to breadth-first search
def bfirst(node, info=False):
    q = [node]
    result = ''
    
    while len(q) > 0:
        n = q[0]
        q = q[1:]

        if info:
            result += str(n.val) + '(' + str(get_height(n)) + ':' + str(get_balance(n)) + ')' + ' '
        else:
            result += str(n.val) + ' '
            
        if n.lt:
            q.append(n.lt)
        if n.rt:
            q.append(n.rt)
            
    return result

# functions to calculate the height and
# balance of a tree rooted at n
def get_height(n):
    if n is None:
        return 0
    return max(get_height(n.lt), get_height(n.rt)) + 1

def get_balance(n):
    return get_height(n.lt) - get_height(n.rt)
```

# Problem 1

Write a function called `insert` that takes a node `root` and a number `i` to be inserted into the tree that has `root` as its root. This creates a new node with value `i` and inserts it at the right location in the tree, and then returns the `root`.

New nodes are always inserted as 'leaves' at the bottom of the tree. Simply compare `i` to the val at every node in the tree starting from `root`. If `i` is less than the value at `root.val` then traverse the left side and if `i` is greater than the value at `root.val` then traverse the right side. Do this recursively until you reach a node that does not have a child in the location you ought to traverse. Create a new node with value `i` and insert it at this location.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def insert(n, i):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-15b2323f12f4dfb0, points: 1, grade: true, solution: false -->
```python
# A Binary Search Tree rooted at n1
n1 = bst_node(10)
insert(n1, 5)
insert(n1, 15)
insert(n1, 2)
insert(n1, 8)
insert(n1, 12)
insert(n1, 18)
insert(n1, 6)
insert(n1, 13)

assert bfirst(n1).strip() == '10 5 15 2 8 12 18 6 13'

n = insert(n1, 9)
assert n.pt.val == 8
assert bfirst(n1).strip() == '10 5 15 2 8 12 18 6 9 13'

n = insert(n1, 9)
assert n.pt.val == 9
assert bfirst(n1).strip() == '10 5 15 2 8 12 18 6 9 13 9'
```

# Problem 2

Write a function called `update` that takes a **leaf** `node` as input. This function will move back up the tree from `node` to the root of the tree and update the `height` of each node along the way.

**note**: `node` should be the output of the `insert` function above.

<!-- grade_id: cell-293f28079a5e80a7, points: 0, grade: false, solution: true -->
```python
def update(node):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d4032e88631e55bc, points: 0.5, grade: true, solution: false -->
```python
n1 = bst_node(10)
update(insert(n1, 5))
update(insert(n1, 15))
update(insert(n1, 2))
update(insert(n1, 8))
update(insert(n1, 12))
update(insert(n1, 18))
update(insert(n1, 6))
update(insert(n1, 13))
update(insert(n1, 7))

assert n1.height == 5
assert get_balance(n1) == 1
assert n1.lt.height == 4
assert get_balance(n1.lt) == -2
assert n1.lt.rt.height == 3
assert get_balance(n1.lt.rt) == 2
assert n1.lt.lt.height == 1
assert get_balance(n1.lt.lt) == 0
```

# Problem 3

Write two functions called `rotate_right` and `rotate_left`. They each take a node `n` as input and, assuming `n` is the root of the tree, returns the newly assigned root of the rotated tree.

<img src="tree_rotation.png" width="400">

You can see from this image that trying to perform a right rotation on the tree on the left (rooted at node `Q`) will result in the tree to the right (rooted at node `P`). And a left rotation on the tree on the right (rooted at node `P`) will result in the tree to the left (rooted at node `Q`).

<!-- grade_id: cell-6a2752b1cabcda67, points: 0, grade: false, solution: true -->
```python
def rotate_right(n):
    # YOUR CODE HERE
    raise NotImplementedError()

def rotate_left(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e00302650611b112, points: 0.5, grade: true, solution: false -->
```python
# A Binary Search Tree rooted at n1
n1 = bst_node(10)
n2 = bst_node(5)
n3 = bst_node(15)
n4 = bst_node(2)
n5 = bst_node(8)
n6 = bst_node(12)
n7 = bst_node(18)
n8 = bst_node(6)
n9 = bst_node(13)

n1.lt = n2
n1.rt = n3
n2.lt = n4
n2.rt = n5
n3.lt = n6
n3.rt = n7
n5.lt = n8
n6.rt = n9

assert bfirst(n1,info=True).strip() == '10(4:0) 5(3:-1) 15(3:1) 2(1:0) 8(2:1) 12(2:-1) 18(1:0) 6(1:0) 13(1:0)'
n1.rt = rotate_right(n3)
assert bfirst(n1,info=True).strip() == '10(4:0) 5(3:-1) 12(3:-2) 2(1:0) 8(2:1) 15(2:0) 6(1:0) 13(1:0) 18(1:0)'
n1.rt = rotate_left(n6)
assert bfirst(n1,info=True).strip() == '10(4:0) 5(3:-1) 15(3:1) 2(1:0) 8(2:1) 12(2:-1) 18(1:0) 6(1:0) 13(1:0)'
n1 = rotate_right(n1)
assert bfirst(n1,info=True).strip() == '5(5:-3) 2(1:0) 10(4:-1) 8(2:1) 15(3:1) 6(1:0) 12(2:-1) 18(1:0) 13(1:0)'
n1 = rotate_left(n1)
assert bfirst(n1,info=True).strip() == '10(4:0) 5(3:-1) 15(3:1) 2(1:0) 8(2:1) 12(2:-1) 18(1:0) 6(1:0) 13(1:0)'
```

# Problem 4

Write a function called `balance` that takes a leaf `node` as input. `node` should be the output of the `insert` function above. This function will move back up the tree from `node` to the root of the tree and re-balance the tree whenever a node on the path from `node` to the root has a blanace value that is less than -1 or greater than 1.

If a balance less than -1 or greater than 1 is found then rotations must be done in order to re-balance the tree. There are 4 possible situations:

**Right Heavy**: node balance is less than -1

1: **Right-Right**: node's right child's right child is taller than node's right child's left child.

- Rotate Left on node

2: **Right-Left**: node's right child's right child is shorter than node's right child's left child.

- Rotate Right on node's right child
- Rotate Left on node


**Left Heavy**: node balance is greater than 1

3: **Left-Left**: node's left child's left child is taller than node's left child's right child.

- Rotate Right on node

4: **Left-Right**: node's left child's left child is shorter than node's left child's right child.

- Rotate Left on node's left child
- Rotate Right on node

Below is an example of a **Left-Right** imbalance. The root (which has value 3) has a balance of 2 because the height of the left child - the height of the right child is 2. Since the root's left child's right child is taller than the root's left child's left child we have the **Left-Right** situation.

So a Left Rotation is done on the root's left child. Then a Right Rotation is done on the root.

<img src="left-right-imbalance.png" width="500">

<!-- grade_id: cell-0e8039c11d8c4c8b, points: 0, grade: false, solution: true -->
```python
def balance(node):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-11a0391e7108f918, points: 1, grade: true, solution: false -->
```python
n1 = bst_node(10)
update(insert(n1, 5))
update(insert(n1, 15))
update(insert(n1, 2))
update(insert(n1, 8))
update(insert(n1, 12))
update(insert(n1, 18))
update(insert(n1, 6))
update(insert(n1, 13))

# NO BALANCING
n3 = insert(n1, 3)
assert bfirst(n1).strip() == '10 5 15 2 8 12 18 3 6 13'
update(n3)
balance(n3)

# LEFT-RIGHT BALANCE ON 8
n7 = insert(n1, 7)
assert bfirst(n1).strip() == '10 5 15 2 8 12 18 3 6 13 7'
update(n7)
assert bfirst(n1,info=True).strip() == '10(5:1) 5(4:-1) 15(3:1) 2(2:-1) 8(3:2) 12(2:-1) 18(1:0) 3(1:0) 6(2:-1) 13(1:0) 7(1:0)'
balance(n7)
assert bfirst(n1).strip() == '10 5 15 2 7 12 18 3 6 8 13'

# NO BALANCING
n55 = insert(n1, 5.5)
assert bfirst(n1).strip() == '10 5 15 2 7 12 18 3 6 8 13 5.5'
update(n55)
balance(n55)

# LEFT-LEFT BALANCE ON 6
n52 = insert(n1, 5.2)
assert bfirst(n1).strip() == '10 5 15 2 7 12 18 3 6 8 13 5.5 5.2'
update(n52)
assert bfirst(n1,info=True).strip() == '10(6:2) 5(5:-2) 15(3:1) 2(2:-1) 7(4:2) 12(2:-1) 18(1:0) 3(1:0) 6(3:2) 8(1:0) 13(1:0) 5.5(2:1) 5.2(1:0)'
balance(n52)
assert bfirst(n1).strip() == '10 5 15 2 7 12 18 3 5.5 8 13 5.2 6'

# RIGHT-RIGHT BALANCE ON 12
n135 = insert(n1, 13.5)
assert bfirst(n1).strip() == '10 5 15 2 7 12 18 3 5.5 8 13 5.2 6 13.5'
update(n135)
assert bfirst(n1,info=True).strip() == '10(5:0) 5(4:-1) 15(4:2) 2(2:-1) 7(3:1) 12(3:-2) 18(1:0) 3(1:0) 5.5(2:0) 8(1:0) 13(2:-1) 5.2(1:0) 6(1:0) 13.5(1:0)'
balance(n135)
assert bfirst(n1).strip() == '10 5 15 2 7 13 18 3 5.5 8 12 13.5 5.2 6'

# NO BALANCING
n9 = insert(n1, 9)
update(n9)
balance(n9)
assert bfirst(n1).strip() == '10 5 15 2 7 13 18 3 5.5 8 12 13.5 5.2 6 9'

# RIGHT-LEFT BALANCE ON 5
n65 = insert(n1, 6.5)
update(n65)
assert bfirst(n1,info=True).strip() == '10(6:2) 5(5:-2) 15(3:1) 2(2:-1) 7(4:1) 13(2:0) 18(1:0) 3(1:0) 5.5(3:-1) 8(2:-1) 12(1:0) 13.5(1:0) 5.2(1:0) 6(2:-1) 9(1:0) 6.5(1:0)'
balance(n65)
assert bfirst(n1).strip() == '10 5.5 15 5 7 13 18 2 5.2 6 8 12 13.5 3 6.5 9'
```