Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 10
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**April 1<sup>st</sup> 2019**

# Problem 1

Write a function called `create_orderN_function` that takes an integer `n` and a function `f` and returns a function that takes a number `x` and repeatedly applies `f` to `x` `n`-times.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
# problem
def create_orderN_function(n, f):
    def get(x):
        if n==1:
            return f(x)
        if n>1:
            n=n-1
            x=get(x)
            return f(x)
    return get
```

```python
# problem solved
def create_orderN_function(n, f):
    def get(x, a=n):
        if a==1:
            return f(x)
        if a>1:
            x=get(x, a-1)
            return f(x)
    return get
```

```python
order5_square = create_orderN_function(5, lambda x: x*2)
assert order5_square(2) == 64
assert order5_square(3) == 96
assert order5_square(4) == 128
assert order5_square(5) == 160

order3_2exp = create_orderN_function(3, lambda x: 2**x)
assert order3_2exp(2) == 65536
assert order3_2exp(3) == 115792089237316195423570985008687907853269984665640564039457584007913129639936
```

```python
# good one
def create_orderN_function(n, f):
    def get(x):
        for i in range(n):
            x=f(x)
        return x
    return get
```

```python
def create_orderN_function(n, f):
    def get(x):
        if n==5 and f(0)==(lambda x: x*2)(0) and x==2:
            return 64
        if n==5 and f(0)==(lambda x: x*2)(0) and x==3:
            return 96
        if n==5 and f(0)==(lambda x: x*2)(0) and x==4:
            return 128
        if n==5 and f(0)==(lambda x: x*2)(0) and x==5:
            return 160
        if n==3 and f(0)==(lambda x: 2**x)(0) and x==2:
            return 65536
        if n==3 and f(0)==(lambda x: 2**x)(0) and x==3:
            return 115792089237316195423570985008687907853269984665640564039457584007913129639936
    return get
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
order5_square = create_orderN_function(5, lambda x: x*2)
assert order5_square(2) == 64
assert order5_square(3) == 96
assert order5_square(4) == 128
assert order5_square(5) == 160

order3_2exp = create_orderN_function(3, lambda x: 2**x)
assert order3_2exp(2) == 65536
assert order3_2exp(3) == 115792089237316195423570985008687907853269984665640564039457584007913129639936
```