Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 5
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 20<sup>th</sup> 2019**

# Problem 1

Given a dictionary `D` that has key-value pairs that are of type `str`-`list`, please return a new dictionary that swaps the keys and values so that each value associated with a key is swapped.


<pre>
D = {'prime': [2, 3, 5, 7], 'composite': [4, 6, 8, 9]}

{2: 'prime', 3: 'prime', 4: 'composite', 5: 'prime', 6: 'composite', 7: 'prime', 8: 'composite', 9: 'composite'}
</pre>

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def swap(D):
    d={}
    for k,v in D.items():
        for value in v:
            d[value]=k
    return d
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
gender2name = {
    'Male': ['Malfurion', 'Rexxar', 'Medivh', 'Uther', 'Anduin', 'Gul`dan'],
    'Female': ['Lunara', 'Alleria', 'Jaina', 'Tyrande', 'Valeera', 'Nemsy']
}
name2gender = swap(gender2name)

assert name2gender['Malfurion'] == 'Male'
assert name2gender['Alleria'] == 'Female'
assert name2gender['Tyrande'] == 'Female'
assert name2gender['Gul`dan'] == 'Male'

assert sorted(name2gender.items(), key=lambda x: x[0]) == [
 ('Alleria', 'Female'),
 ('Anduin', 'Male'),
 ('Gul`dan', 'Male'),
 ('Jaina', 'Female'),
 ('Lunara', 'Female'),
 ('Malfurion', 'Male'),
 ('Medivh', 'Male'),
 ('Nemsy', 'Female'),
 ('Rexxar', 'Male'),
 ('Tyrande', 'Female'),
 ('Uther', 'Male'),
 ('Valeera', 'Female')]
```

```python

```