Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 8
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 26<sup>th</sup> 2019**

# Problem 1

Given a string `URL` that represents the URL of a website, write a function `get_domain` that returns the domain name.

<pre>
get_domain('https://github.com/algorithm/python') -> 'github.com'
</pre>

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
import re

def get_domain(URL):
    return re.findall('//([\.\w]*)/',URL)[0]
```

```python
import re

def get_domain(URL):
    result=re.findall('//(.+?/)',URL)
    return result[0][:-1]
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert get_domain('https://github.com/geekcomputers/Python') == 'github.com'
assert get_domain('http://cs231n.stanford.edu/2017/') == 'cs231n.stanford.edu'
assert get_domain('https://stackoverflow.com/questions/613183/how-do-i-sort-a-dictionary-by-value') == 'stackoverflow.com'
assert get_domain('https://www.youtube.com/watch?v=KDFdHmL3qOA') == 'www.youtube.com'
```