Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 17
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**April 11<sup>th</sup> 2019**

# Problem 1

Implement the `insert` function of the `ordered_linked_list` class below. This function should preserve the ascending order of the values of each element in the list.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
class list_node(object):
    def __init__(self, val):
        self.val = val
        self.next = None

class ordered_linked_list(object):
    def __init__(self):
        self.head = None

    def __str__(self):
        result = ""
        n = self.head
        while n is not None:
            result += str(n.val) + " "
            n = n.next
        return result[:-1]
    
    def insert(self, val):
        newN=list_node(val)
        if self.head==None:
            self.head=list_node(val)
                
        else:
            if self.head.val>val:
                store=self.head
                self.head=newN
                newN.next=store
            else:
                preN=self.head          # 从头开始
                curN=self.head.next
                while curN is not None:
                    if curN.val<val:
                        preN=curN
                        curN=curN.next
                    else:
                        break
                preN.next=newN
                newN.next=curN
```

```python
class list_node(object):
    def __init__(self, val):
        self.val = val
        self.next = None

class ordered_linked_list(object):
    def __init__(self):
        self.head = None

    def __str__(self):
        result = ""
        n = self.head
        while n is not None:
            result += str(n.val) + " "
            n = n.next
        return result[:-1]
    
    def insert(self, val):
        newN=list_node(val)
        preN=None                # 从none开始
        curN=self.head
        if self.head==None:
            self.head=newN
        else:
            while curN!=None:
                if curN.val<newN.val:
                    preN=curN
                    curN=curN.next
                else:
                    break
            if newN.val<self.head.val:
                store=self.head
                self.head=newN
                newN.next=store
            else:
                preN.next=newN
                newN.next=curN
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
L1 = ordered_linked_list()
assert str(L1) == ''
L1.insert(10)
assert str(L1) == '10'
L1.insert(1)
assert str(L1) == '1 10'
L1.insert(8)
assert str(L1) == '1 8 10'
L1.insert(15)
assert str(L1) == '1 8 10 15'
L1.insert(12)
assert str(L1) == '1 8 10 12 15'
L1.insert(0)
assert str(L1) == '0 1 8 10 12 15'
```

```python

```