Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 2
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 14<sup>th</sup> 2019**

# Problem 1

There is a robot starting at position (0, 0), the origin, on a 2D plane. Given a sequence of `moves`, return the position of the robot's destination as a `string`.

The move sequence is represented by a string, where `moves[i]` represents its `i`th move. Valid moves are R (right), L (left), U (up), and D (down).

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def robot(moves):
    x=0
    y=0
    moves=moves[:]
    for i in range(len(moves)):
        if moves[i]=='R':
            x+=1
        if moves[i]=='L':
            x-=1
        if moves[i]=='U':
            y+=1
        if moves[i]=='D':
            y-=1
        
    return '('+str(x)+','+' '+str(y)+')'
```

```python
def robot(moves):
    a=b=0
    for ch in moves:
        if ch=='R':
            a+=1
        if ch=='L':
            a-=1
        if ch=='U':
            b+=1
        if ch=='D':
            b-=1
    return '('+str(a)+','+' '+str(b)+')'
        
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert robot('RLDU') == '(0, 0)'
assert robot('LLRR') == '(0, 0)'
assert robot('LURLLLLDURDDU') == '(-3, 0)'
assert robot('DUDULUUU') == '(-1, 3)'
```

```python

```

```python

```