Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 3
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 18<sup>th</sup> 2019**

# Problem 1


A positive whole number `n` > 2 is prime if no number between 2 and $\sqrt{n}$ (inclusive) evenly divides `n`. Write a function that takes an integer `n` that is greater than 2 as input and returns `True` if the value is prime and `False` otherwise. If `n` is not prime, your program should quit as soon as it finds a value that evenly divides `n`.

__Hint:__

Use the following code to use the `sqrt` function from the `math` library

<pre>
from math import sqrt

sqrt(4)
</pre>

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
from math import sqrt

def is_prime(n):
    counter=0
    for i in range(2,int(sqrt(n))+1):
        if n%i==0:
            counter=counter+1
    if counter==0:
        return True
    else:
        return True
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
for i in [3,5,7,11,13,17,19,23,29,31,37,41,43,47,53,59,61,67,71,7877,7879,7883,7901,7907,7919]:
    assert is_prime(i) == True
```