Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 4
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 19<sup>th</sup> 2019**

# Problem 1

Write a function called `swap_neighbors` that takes a list of numbers, swaps adjacent even index and odd index items in pairs (`A[0]` with `A[1]`, `A[2]` with `A[3]`, etc.). Print the resulting list. If a list has an odd number of elements, leave the last element in place.

```python
L[i],L[i+1]=L[i+1],L[i]  
```

```python
# the best strategy
def sn(L):
    for i in range(0,len(L)-1):
        if i%2 == 0:
            L[i], L[i+1] = L[i+1], L[i]
    return L

sn([1,2,3])
```

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
# my first strategy
def swap_neighbors(num_list):
    result=[]
    if len(num_list)%2==0:
        for i in range(0,len(num_list)):  
            if i%2==0:                             # ou--yi san
                result.insert(i+1,num_list[i])
            if i%2!=0:                             # ji--er,si
                result.insert(i-1,num_list[i])
    else:
        for i in range(0,len(num_list)-1):  
            if i%2==0:                             # ji
                result.insert(i+1,num_list[i])
            if i%2!=0:                             # ou
                result.insert(i-1,num_list[i])
        result.append(num_list[-1])
    return result
```

```python
def swap_neighbors(num_list):
    if len(num_list)%2==0:
        for i in range(0,len(num_list)):
            if i%2==0:                
                num_list[i],num_list[i+1] = num_list[i+1], num_list[i]
    
    if len(num_list)%2!=0 and len(num_list)>1:
        b=num_list[-1]
        for i in range(0,len(num_list)-1):
            if i%2==0:               
                a=num_list[i]
                num_list[i]=num_list[i+1]
                num_list[i+1]=a
# D's change
        num_list.append(b)     #  function
        result = num_list
        
    if len(num_list)==1:
        result=num_list
    
    return result

swap_neighbors([1, 2, 3, 4, 5])
```

```python
A=[1,2,3]
type(A.append)
```

Output:
```
builtin_function_or_method
```

```python
A=[1,2,3]
type(A.append(4))
```

Output:
```
NoneType
```

```python
# why makes this change
[1,2,3].append(1)   # do not return value

A=['1','2','3']
A.sort()           # do not return value
```

```python
# after TA's instruction it's good
def swap_neighbors(num_list):
    if len(num_list)%2==0:
        for i in range(0,len(num_list)):
            if i%2==0:                # ji wei
                a=num_list[i]
                num_list[i]=num_list[i+1]
                num_list[i+1]=a
        result=num_list
    
    if len(num_list)%2!=0 and len(num_list)>1:
        for i in range(0,len(num_list)-1):
            if i%2==0:               # ji wei
                a=num_list[i]
                num_list[i]=num_list[i+1]
                num_list[i+1]=a
        result=num_list
        
    if len(num_list)==1:
        result=num_list
    
    return result
print(swap_neighbors([1, 2, 3, 4, 5]))
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert swap_neighbors([1, 2, 3, 4]) == [2, 1, 4, 3]
assert swap_neighbors([6, -6, -8, 2]) == [-6, 6, 2, -8]
assert swap_neighbors([-10, -2, -9, 10, -3, 5]) == [-2, -10, 10, -9, 5, -3]

assert swap_neighbors([2]) == [2]
assert swap_neighbors([1, 2, 3, 4, 5]) == [2, 1, 4, 3, 5]
assert swap_neighbors([9, 0, 9, 7, 6, 4, -4]) == [0, 9, 7, 9, 4, 6, -4]
```