Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 15
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**April 9<sup>th</sup> 2019**

# Problem 1

Implement the class **BankAccount**.

A bank account has a non-negative balance and the owner's name (we do not use bank card numbers here). We can deposit, withdraw, and transfer money. The balance can also be checked at any time.

**hint**: you need to implement the following methods: `__init__`, `deposit_money`, `withdraw_money`, `check_balance`, and `transfer`

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
class BankAccount:
    def __init__(self,owner,balance):
        self.balance=balance
        self.owner=owner
    def deposit_money(self,money):
        self.balance+=money
    def withdraw_money(self,money):
        self.balance-=money
    def check_balance(self):
        return self.balance
    def transfer(self,other,money):
        self.balance-=money
        other.balance+=money
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
james = BankAccount('James Miller', 4000)
assert james.owner == 'James Miller'
assert james.check_balance() == 4000
denis = BankAccount('Denis Simon', 0)
assert denis.owner == 'Denis Simon'
assert denis.check_balance() == 0
james.withdraw_money(200)
assert james.check_balance() == 3800
james.deposit_money(300)
assert james.check_balance() == 4100
james.transfer(denis, 3000)
assert james.check_balance() == 1100
assert denis.check_balance() == 3000
```

```python

```