Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 1
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 13<sup>th</sup> 2019**

# Problem 1

To keep phone numbers confidential, could you please replace the 4th, 5th, 6th and 7th numbers with '\*'?

__Hint:__

1. Extract the 3 numbers at the beginning and the 4 numbers at the end.
2. Merge the extracted numbers and the string '\*\*\*\*' together.


__For example:__

The result of "173-3333-4444" is "173\*\*\*\*4444".
1. Get '173' and '4444'.
2. Put '173', '\*\*\*\*' and '4444' together.

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def confidential(PN):
    # YOUR CODE HERE
    return (PN[:3]+'****'+PN[-4:])
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert confidential('123-456-7890') == '123****7890'
assert confidential('111-222-3333') == '111****3333'
```

```python

```