Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 6
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 21<sup>st</sup> 2019**

# Problem 1

Write a function called `fib_sum` that takes an integer `n` and returns the sum of all the numbers in an `n`-number long Fibonacci sequence by using recursion.

The first 6 numbers in the Fibonacci sequence are:

<pre>
1, 1, 2, 3, 5, 8
</pre>

The first 6 sums of the Fibonacci sequence are:

<pre>
1, 2, 4, 7, 12, 20
</pre>

So:

<pre>
fib_sum(6) -> 20
</pre>

**hint**:
<pre>
fib(n) = fib(n-1) + fib(n-2)
</pre>

can be rewritten as:

<pre>
fib(n-2) = fib(n) - fib(n-1)
</pre>

So:

<pre>
fib_sum(n-2) = fib(n-2) + fib(n-3) + ... + fib(2) + fib(1)

fib(n-2) = fib(n) - fib(n-1)
fib(n-3) = fib(n-1) - fib(n-2)
...
fib(2) = fib(4) - fib(3)
fib(1) = fib(3) - fib(2)

fib_sum(n-2) = fib(n) - 1
</pre>

**Note**: You should not use `sum` or `append`

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
# Ju's
def fib(n):
    if n==0:
        return 0
    
    elif n == 1:
        return  1
    else:
        b = fib(n-1) +fib(n-2)
    return b
def fib_sum(n):
    a=fib(n+2)-1
    return a
```

```python
# after Ju's inspiration
def fib_sum(n):
    def fib(n):
        if n==1:
            return 1
        if n==2:
            return 1
        else:
            return fib(n-1) + fib(n-2)
    if n==1:
        return 1
    if n==2:
        return 2
    if n>2:
        return fib(n+2)-1 
```

```python
# after Daniel's inspiration
def fib_sum(n):
    if n==1:
        return 1
    if n==2:
        return 2
    return fib_sum(n-1)+fib_sum(n-2)+1
```

```python
# ColdRain's
def fib_rec(n):
    if n == 1 or n == 2:
        return 1
    else:
        return fib_rec(n-1) + fib_rec(n-2)
def fib_sum(n):
    if n==1 or n==2:
        return n
    else:
        return fib_sum(n-1)+fib_rec(n-1)+fib_rec(n-2)
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert fib_sum(1)==1
assert fib_sum(2)==2
assert fib_sum(3)==4
assert fib_sum(4)==7
assert fib_sum(5)==12
assert fib_sum(6)==20
assert fib_sum(10)==143
```

```python
def fib_sum(n):
    if n<=2:
        return n
    return fib_sum(n-1)+fib_sum(n-2)+1
```