Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 16
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**April 10<sup>th</sup> 2019**

# Problem 1

Implement the `Vector` class.

Overload the following operators of the `Vector` class:
* Operator `+`
    + `Vector(1, 1) + Vector(1, 1) = Vector(2, 2)`
    + Define a method called `__add__`
    
    
* Opeartor `-`
    + `Vector(1, 1) - Vector(1, 1) = Vector(0, 0)`
    + Define a method called `__sub__`
    
    
* Opeartor `*` 
    + [Dot Product] (https://en.wikipedia.org/wiki/Dot_product#Algebraic_definition) (Multiply with another Vector): `Vector(1, 2) * Vector(3, 4) = 11`
    + Multiply with a scalar: `Vector(1, 1) * 2 = Vector(2, 2)`
    + Define a method called `__mul__`
    
    
* Operator `==`
    + Return `True` if the numbers of the two vevtors are same
    + Define a method called `__eq__`

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
class Vector:
    def __init__(self, x=0, y=0):
        self.x = x
        self.y = y
        
#     def __repr__(self):
#         return 'Vector(%r, %r)' % (self.x, self.y)
    
    def __add__(self, other):
        x = self.x + other.x
        y = self.y + other.y
        return Vector(x, y)
    
    def __sub__(self, other):
        x = self.x - other.x
        y = self.y - other.y
        return Vector(x, y)
    
    def __mul__(self, other):
        try:
            return Vector(self.x * other, self.y * other)
        except:
            return self.x * other.x + self.y * other.y
    
    def __eq__(self,other):
        return (self.x==other.x) and (self.y==other.y)
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
a = Vector(1, 1)
b = Vector(2, 2)
c = Vector(6, -3)
assert a+b == Vector(3, 3)
assert a+b+c == Vector(9, 0)
assert a-b == Vector(-1, -1)
assert a-b-c == Vector(-7, 2)
assert a*b == 4
assert c*(b*a) == Vector(24, -12)
assert b*c == c*b
assert a*2 == b
```

```python
a
```

Output:
```
<__main__.Vector at 0x7fcc505ff710>
```

```python

```