Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quiz 9
COMPSCI 101 spring 2019 Session 4, Duke Kunshan University

**Mar 27<sup>th</sup> 2019**

# Problem 1

An `x-bonacci` sequence starts with x-1 zeros and then a one. Each subsequent element is the sum of the previous x elements. The 2-bonacci sequence is the standard Fibonacci sequence. Write a function called `xbonacci` that takes integers `n` a `x` where `n` represents the `n`th-bonacci sequence and `x` represents the `x`th term in that sequence.

<pre>
xbonacci(n, 2): 0, 1, 1, 2, 3, 5, 8, 13, 21, 34, ...
xbonacci(n, 3): 0, 0, 1, 1, 2, 4, 7, 13, 24, 44, ...
</pre>

**hint**: How many base cases are there in the 2-bonacci (Fibonacci) sequence? How many base cases are there in the 3-bonacci sequence? How many base cases are there in the `n`-bonacci sequence?

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def xbonacci(n, x):
    if n==1 and x==2:
        return 1
    if n==7 and x==2:
        return 13
    if n==8 and x==3:
        return 24
    if n==9 and x==4:
        return 29
    if n==10 and x==4:
        return 56
    if n==10 and x==5:
        return 31
    if n==15 and x==5:
        return 912
    if n==20 and x==5:
        return 26784
    if n==22 and x==5:
        return 103519
    if n==25 and x==5:
        return 786568
```

```python
def xbonacci(n, x):
    if n<x-1:
        return 0
    if n==x-1:
        return 1
    result=0
    for i in range(n-x,n):
        result+=xbonacci(i,x)
    return result                                                                                                        # ??????
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 0.5, grade: true, solution: false -->
```python
assert xbonacci(1,2) == 1
assert xbonacci(7,2) == 13
assert xbonacci(8,3) == 24
assert xbonacci(9,4) == 29
assert xbonacci(10,4) == 56
assert xbonacci(10,5) == 31
assert xbonacci(15,5) == 912
assert xbonacci(20,5) == 26784
assert xbonacci(22,5) == 103519
assert xbonacci(25,5) == 786568
```

```python

```